# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 6

```python
%matplotlib inline
import numpy as np
import networkx as nx
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
from numpy.testing import assert_almost_equal
import math
```

# First Course: chap 3.3 Friendship Paradox

1. Maak een netwerkx netwerk van Figuur 3.5.
2. Waarom is de kans dat je Tom vindt met de "random friend of picked person" methode $\frac{5}{21}$. Geef een heel helder antwoord.
3. Programmeer de functie `kans_op_knoop(G,knoop)` die voor elk netwerk, voor elke knoop de kans berekent dat je die knoop trekt met de "random friend of picked person" methode.
4. Schrijf een test die checked dat die kansen optellen tot 1.
5. **Advanced** In de een na laatste paragraaf van 3.3 op blz 76 wordt een andere manier om die kansen te berekenen gemeld, die dezelfde uitkomst geeft. Programmeer die ook, en vergelijk de uitkomsten.
    

<!-- grade_id: FPa, points: 0, grade: false, solution: true -->
```python
Fig35=nx.Graph()
nodes = ['mary', 'nancy', 'tom', 'john', 'pam', 'bob', 'tara']
blue_node = ['tom']
Fig35.add_nodes_from(nodes)

Fig35.add_edges_from([('mary', 'nancy'), ('nancy', 'tom'), ('nancy', 'bob'), ('bob', 'tara'), ('tom', 'pam'),
                     ('john', 'pam'), ('tom', 'john'),('tom', 'bob')])
for node in Fig35: 
    if not node in blue_node:
        Fig35.nodes[node]['B'] = False
        Fig35.nodes[node]['R'] = True
    else:
        Fig35.nodes[node]['B'] = True 
        Fig35.nodes[node]['R'] = False

color_map = []
for node in Fig35:
    if Fig35.nodes[node]['R'] == False:
        color_map.append('blue')
    else: 
        color_map.append('red')      
nx.draw(Fig35, node_color=color_map, with_labels=True)
plt.show()
        

def kans_op_knoop(G,knoop):
    #WRITE YOUR CODE HERE
    kans = []
    for i in G:
        if G.has_edge(knoop,i):
            kans.append((1/G.degree(i)) * (1/nx.number_of_nodes(G)))
        else:
            kans.append(0)
    return sum(kans)

 

H = nx.karate_club_graph() 
# test
sum(kans_op_knoop(H,x) for x in H)
#kans_op_knoop(H,1)
```

<!-- grade_id: FPt, points: 1, grade: true, solution: false -->
```python
assert isinstance(Fig35, nx.Graph)
H= nx.karate_club_graph()
assert 0 <= kans_op_knoop(H,1)<= 1

```

# First course

Maak de sommen 3.11-3.24.

<!-- grade_id: fc, points: 14, grade: true, solution: true -->
3.11 
1) Node 5, node 2 en node 3
2) Node 2
3) Node 2

3.12
Decgraph heeft een centrale keten van knopen met één of een paar grote knooppunten. De knopen aan deze knooppunten hebben vaak zelf geen buren, waardoor er een soort stervorm ontstaat.

3.13
1) Nee, ik denk dat dit eerder normaal verdeeld is. De meeste mensen hebben waarschijnlijk een schoenmaat tussen 39 en 43. Daarbuiten zal de grafiek steeds meer aflopen.
2) Ja, er zijn heel veel gezinnen met een laag / gemiddeld inkomen en maar heel weinig gezinnen met een heel hoog inkomen.
3) Ja, er zijn heel veel account met maar een paar vrienden / volgers, maar er zijn maar weinig accounts met heel veel vrienden / volgers.
4) Nee, deze afstand hangt waarschijnlijk af van de gelijkenis van twee wikipediapagina's.

3.14
De kans zou niet groot zijn, maar het zou mogelijk zijn volgens deze verdeling.

3.15
1) Ongeveer 1*10^8 paginas
2) Ongeveer 1*10^7 paginas
3) Ongeveer 1*10^4 paginas

3.16
De maximal degree in het onderzoek is 7. Dit betekent dat er geen mensen in het onderzoek waren met meer partners.

<!-- grade_id: fc324, points: 6, grade: true, solution: true -->
```python
WRITE YOUR CODE HERE
```

```python

```