# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 6

```python
%matplotlib inline
import numpy as np
import networkx as nx
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
from numpy.testing import assert_almost_equal
import math
```

# First Course: chap 3.3 Friendship Paradox

1. Maak een netwerkx netwerk van Figuur 3.5.
2. Waarom is de kans dat je Tom vindt met de "random friend of picked person" methode $\frac{5}{21}$. Geef een heel helder antwoord.
3. Programmeer de functie `kans_op_knoop(G,knoop)` die voor elk netwerk, voor elke knoop de kans berekent dat je die knoop trekt met de "random friend of picked person" methode.
4. Schrijf een test die checked dat die kansen optellen tot 1.
5. **Advanced** In de een na laatste paragraaf van 3.3 op blz 76 wordt een andere manier om die kansen te berekenen gemeld, die dezelfde uitkomst geeft. Programmeer die ook, en vergelijk de uitkomsten.
    

<!-- grade_id: FPa, points: 0, grade: false, solution: true -->
```python
Fig35=nx.Graph()
nodes = ['Mary', 'Nancy', 'Tom', 'John','Pam', 'Bob', 'Tara']
edges = [('Mary', 'Nancy'), ('Nancy', 'Tom'), ('Nancy', 'Bob'), ('Tom', 'John'), 
         ('Tom', 'Pam'), ('Tom', 'Bob'), ('John', 'Pam'), ('Bob', 'Tara')]
Fig35.add_nodes_from(nodes)
Fig35.add_edges_from(edges)
"""
De 5 / 21 kans is als volgt opgebouwd:
Voor elke node berekenen we de kans dat deze node naar Tom leidt:
Dit is voor elke node: 1 / len(G.nodes) * 1 / G.degree(node)
Als een node geen verbinding heeft met Tom, is deze berekening 1 / len(G.nodes) * 0
De berekening wordt in dit geval: 1/7*0+1/7*0+1/7*1/3+1/7*1/3+1/7*1/2+1/7*1/2 = 5/21
De som nemen van al deze kansen komt dus uit op 5 / 21
""" 

def kans_op_knoop(G, knoop):
    chances = []
    connected_with_node = G.neighbors(knoop)
    for node in connected_with_node:
        chance = (1 / len(G)) * (1 / G.degree(node))
        chances.append(chance)
    return sum(chances) 

def kans_op_knoop_alternatief(G, knoop):
    chances = []
    connected_with_node = G.neighbors(knoop)
    for node in connected_with_node:
        chance = 1 / G.degree(node)
        chances.append(chance)
    return sum(chances) / G.number_of_nodes()

H= nx.karate_club_graph() 
# test
sum(kans_op_knoop(H,x) for x in H)
#kans_op_knoop(H,1)
```

Output:
```
1.0
```

<!-- grade_id: FPt, points: 1, grade: true, solution: false -->
```python
assert isinstance(Fig35, nx.Graph)
H= nx.karate_club_graph()
assert 0 <= kans_op_knoop(H,1)<= 1

```

# First course

Maak de sommen 3.11-3.24.

<!-- grade_id: fc, points: 14, grade: true, solution: true -->
3.11.a: Nodes 2 and 3. \
3.11.b: Nodes 2 and 5. \
3.11.c: Node 2. \
3.12: The graph would look like a star (A hub with a lot of connections and many nodes connected to that hub) \
3.13.1: No heavy-tailed distribution, because not a lot of people have a shoe size that is super above average. \
3.13.2: Heavy-tailed distibution, because a small amount of households has a way larger income than the median income. \
3.13.3: Heavy-tailed distribution, because a small amount of people have a degree that is way higher than the median degree (celebrities/politicians/etc) \
3.13.4: No heavy-tailed distribution, because this pairwise distance is usually really average. \
3.14: It would still be surprising, but less surprising than it would be without a heavy-tailed distribution.
A heavy-tailed distribution means that there is a significant amount of outliers, but not necessarily higher or lower outliers. \
3.15.1: Approximately 100.000.000 pages \
3.15.2: Approximately 5.000.000 pages \
3.15.3: Approximately 10.000 pages \
3.16: ? max degree ? what it means ? if they have the same degree distribution ? why

<!-- grade_id: fc324, points: 6, grade: true, solution: true -->
# 3.17
filenameg='C:/Users/jesse/Desktop/NetworkScience/FirstCourseNetworkScience-master/datasets/openflights/openflights_world.edges.gz'
G= nx.Graph()
G = nx.read_edgelist(filenameg)

#3.17.1
routes = {node: G.degree(node) for node in G}
average_num_routes = sum(routes.values()) / len(routes)
# So the average number of routes served by each airport in 11.71.
#3.17.2
routes_top_5 = sorted(routes, key = routes.get, reverse = True)[:5]
# So the top five airports in terms of number of routes are ['AMS', 'FRA', 'CDG', 'IST', 'ATL'].
#3.17.3
one_route_airports = [x[0] for x in routes.items() if x[1] == 1]
amount_one_route_airports = len(one_route_airports)
# So there are 709 airports that only serve one route.
#3.17.4
closeness_centrality = nx.closeness_centrality(G)
highest_closeness_centrality = max(closeness_centrality, key = closeness_centrality.get)
# So the airport with the highest closeness centrality is FRA
#3.17.5
betweenness_centrality = nx.betweenness_centrality(G)
highest_betweenness_centrality = max(betweenness_centrality, key = betweenness_centrality.get)
# So the airport with the highest betweenness centrality is CDG
#3.17.6
def heterogeneity_undirected(G):
    degrees_only = [degree for node, degree in G.degree()]
    degrees_squared = [degree**2 for degree in degrees_only]
    degrees_squared_average = sum(degrees_squared) / len(degrees_squared)
    degrees_average = sum(degrees_only) / len(degrees_only)
    degrees_average_squared = degrees_average**2
    heterogeneity = degrees_squared_average / degrees_average_squared
    return heterogeneity
# So the heterogeneity of this network is 5.51

```python
#3.18
filenameh='../../../data/Week6/enwiki_math.edges.gz'
H= nx.read_edgelist(filenameh)
in_degrees = {node: H.in_degree(node) for node in H}
out_degrees = {node: H.out_degree(node) for node in H}
#3.18.1
average_in_degree = sum(in_degrees.values()) / len(in_degrees)
average_out_degree = sum(out_degrees.values()) / len(out_degrees)
# The average in- and out-degree are equal (both 12.75). This means that the page has as many in-links as out-links.
#3.18.2
highest_in_degree = max(in_degrees, key = in_degrees.get)
# So node '1152126' has the highest in_degree
#3.18.3
highest_out_degree = max(out_degrees, key = out_degrees.get)
# So node '47738065' has the highest out_degree
#3.18.4
# The maximum in-degree is higher than the maximum out-degree. I won't expect this to be the case for many other web graphs,
# since mathematics is very general and thus many current theories and ways things work relate back to mathematics. Other,
# more specific pages will probably have a higher out-degree than in-degree, since not a lot of things relate back to
# specific pages, but a specific page is probably based on a lot of general pages. 
#3.18.5
def heterogeneity_directed_in(G):
    in_degrees_only = [in_degree for node, in_degree in G.in_degree()]
    in_degrees_squared = [in_degree**2 for in_degree in in_degrees_only]
    in_degrees_squared_average = sum(in_degrees_squared) / len(in_degrees_squared)
    in_degrees_average = sum(in_degrees_only) / len(in_degrees_only)
    in_degrees_average_squared = in_degrees_average**2
    in_heterogeneity = in_degrees_squared_average / in_degrees_average_squared
    return in_heterogeneity
# So the in_heterogeneity of this network is 38.21
def heterogeneity_directed_out(G):
    out_degrees_only = [out_degree for node, out_degree in G.out_degree()]
    out_degrees_squared = [out_degree**2 for out_degree in out_degrees_only]
    out_degrees_squared_average = sum(out_degrees_squared) / len(out_degrees_squared)
    out_degrees_average = sum(out_degrees_only) / len(out_degrees_only)
    out_degrees_average_squared = out_degrees_average**2
    out_heterogeneity = out_degrees_squared_average / out_degrees_average_squared
    return out_heterogeneity
# So the out_heterogeneity of this network is 6.62

#3.19
def average_degree_neighbors(G, node):
    neighbors = G.neighbors(node)
    neighbor_degrees = [G.degree(neighbor) for neighbor in neighbors]
    average_degree_neighbors = sum(neighbor_degrees) / len(neighbor_degrees)
    return average_degree_neighbors

#3.20
# Yes, such networks would be assortative, meaning that nodes with similar degrees tend to be connected to each other
# (high degrees to high degrees, low degrees to low degrees)

#3.21.1
# Networks with heavy-tailed degree distributions are more vulnerable to targeted attacks, since there is a lesser amount of
# nodes with a high degree (these high degrees are extremely high, but there are not many nodes with such a high degree).
# An attack on one of these nodes would be very damaging for a network.
#3.21.2
# A grid-like network will have a more steady and equal degree distribution, making it less vulnerable for attacks.
# A targeted attack could still be more damaging than a random attack, since a node in a grid-like network
# can play a big role in connected one node to another, since degrees will not be very high in such a network, meaning there
# can be some crucial nodes that are the only ones forming a path between one and another node.
#3.22
# d, nodes with high betweenness centrality. A node with a high betweenness centrality is important for other nodes,
# since it lies in a lot of shortests paths between two nodes. If such a node is removed, a lot of shortest paths
# will become longer, thus increasing the average path length.
#3.23
# The one with a lower clustering coefficient would be a better target. Imagine node A and node B being the nodes
# this assignment is about. node A and B both have degree 10, but every neighbor of node A is connected in some way
# and none of the neighbors of node B are connected to each other, they are only connected to node B itself.
# If node A is removed, the friend nodes will only disconnect from A 
# and will still be reachable from each other friend node of A.
# If node B is removed, the friend nodes will disconnect from B and be disconnected from each other, forming 10 new subnetworks.
# So removing the node with a lower clustering coefficient will disrupt the network more.
```

Output:
```
Cell In[4], line 4
      2 filenameh='../../../data/Week6/enwiki_math.edges.gz'
      3 H= nx.read_edgelist(filenameh)
----> 4 in_degrees = {node: H.in_degree(node) for node in H}
      5 out_degrees = {node: H.out_degree(node) for node in H}
      6 #3.18.1

Cell In[4], line 4, in <dictcomp>(.0)
      2 filenameh='../../../data/Week6/enwiki_math.edges.gz'
      3 H= nx.read_edgelist(filenameh)
----> 4 in_degrees = {node: H.in_degree(node) for node in H}
      5 out_degrees = {node: H.out_degree(node) for node in H}
      6 #3.18.1

AttributeError: 'Graph' object has no attribute 'in_degree'
```

```python

```