# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 6

```python
%matplotlib inline
import numpy as np
import networkx as nx
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
from numpy.testing import assert_almost_equal
import math
```

# First Course: chap 3.3 Friendship Paradox

1. Maak een netwerkx netwerk van Figuur 3.5.
2. Waarom is de kans dat je Tom vindt met de "random friend of picked person" methode $\frac{5}{21}$. Geef een heel helder antwoord.
3. Programmeer de functie `kans_op_knoop(G,knoop)` die voor elk netwerk, voor elke knoop de kans berekent dat je die knoop trekt met de "random friend of picked person" methode.
4. Schrijf een test die checked dat die kansen optellen tot 1.
5. **Advanced** In de een na laatste paragraaf van 3.3 op blz 76 wordt een andere manier om die kansen te berekenen gemeld, die dezelfde uitkomst geeft. Programmeer die ook, en vergelijk de uitkomsten.
    

<!-- grade_id: FPa, points: 0, grade: false, solution: true -->
```python
Fig35=nx.Graph()
nodes=['Tom','John','Pam', 'Mary', 'Nancy', 'Bob', 'Tara']
edges=[('Tom', 'John'), ('John', 'Pam'),('Pam', 'Tom'), ('Tom', 'Nancy'), ('Nancy', 'Mary'), ('Nancy', 'Bob'), ('Bob', 'Tara'), ('Bob', 'Tom')]
Fig35.add_nodes_from(nodes)
Fig35.add_edges_from(edges)

# Als je voor elke knoop die buren zijn van Tom: 1/|N| * 1/degree(n) doet en dat optelt kom je uit op 5/21

def kans_op_knoop(G,knoop):
    aantalN = G.number_of_nodes()
    buren = G.neighbors(knoop)
    totale_kans = 0
    for x in buren: 
        kans = 1/aantalN * 1/G.degree(x)
        totale_kans += kans
    return totale_kans 

H = nx.karate_club_graph() 
# test
sum(kans_op_knoop(H,x) for x in H)
# kans_op_knoop(H,1)
kans_op_knoop(Fig35, 'Tom')
```

Output:
```
0.23809523809523808
```

<!-- grade_id: FPt, points: 1, grade: true, solution: false -->
```python
assert isinstance(Fig35, nx.Graph)
H= nx.karate_club_graph()
assert 0 <= kans_op_knoop(H,1)<= 1

```

# First course

Maak de sommen 3.11-3.24.

3.11: 1. 2,35, 2.2, 3. 4 <br>
3.12: ? <br>
3.13: 1. Nee, want kleine schoenmaten heeft bijna niemand en hetzelfde geldt voor grote schoenmaten. <br>
      2. Ja, want hoe rijker iemand is, hoe minder huishoudens ook zo rijk of rijker zijn, over het algemeen zijn er meer arme mensen in de wereld <br> 
      3. Ja, veel mensen hebben weinig volgers en maar een paar hebben heel veel volgers. <br>
      4. Ik weet niet wat dat is <br>
3.14: Ja <br> 
3.15: 1. 1^08, 2. 1^06, 3. 10.000 <br>
3.16: Maximum degree is degene die met de meeste mensen heeft geslapen. Als het netwerk wordt gesplitst tussen mannen en vrouwen kan het de degree ditrubution anders worden omdat mannnen veel meer rond slapen dan vrouwen of andersom. <br> 
3.17: 1. 10.19, 2.DEN, ATL, ORD, DFW, MSP 3. 156, 4. DEN, 5. ANC


<!-- grade_id: fc, points: 14, grade: true, solution: true -->
```python

fname='../../../data/Week6/openflights_usa.edges.gz'
G= nx.Graph()
G = nx.read_edgelist(fname)# nx.draw_networkx(G)
```

```python
# What is the avarage number of routes served by each airport in this network?
totaal = 0 
for x in G: 
    totaal += G.degree(x)
print(totaal / G.number_of_nodes())

# What are the top five airports in terms of number of routes?
highest = 0 
lijst = [G.degree(x) for x in G]
lijst.sort(reverse=True)
# print(lijst)
for x in G: 
    if G.degree(x) == 153 or G.degree(x) == 149 or G.degree(x) == 148 or G.degree(x) ==138 or G.degree(x) ==117: 
        print(x)

# How many airports in this network serve only a single rourte?
counter = 0 
for x in G:
    if G.degree(x) == 1: 
        counter += 1
print(counter)

# Which airport has the highest closeness centrality?
highest1 = 0
highestv = 0
for x in G: 
    if nx.closeness_centrality(G, x) > highest1: 
        highest1 = nx.closeness_centrality(G, x)
        highestv = x
print(highestv)

# Which airport has the highest betweenness centrality?
dictd = nx.betweenness_centrality(G)
highest2= 0
highestb = 0
for x in dictd: 
    if dictd[x] >= highest2: 
        highest2 = dictd[x]
        highestb = x
print(highestb)
```

Output:
```
10.186813186813186
DEN
ATL
ORD
DFW
MSP
156
DEN
ANC
```

<!-- grade_id: fc324, points: 6, grade: true, solution: true -->
```python
WRITE YOUR CODE HERE
```

Output:
```
Cell In[6], line 1
    WRITE YOUR CODE HERE
          ^
SyntaxError: invalid syntax
```

```python

```