# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 6

```python
%matplotlib inline
import numpy as np
import networkx as nx
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
from numpy.testing import assert_almost_equal
import math
```

# First Course: chap 3.3 Friendship Paradox

1. Maak een netwerkx netwerk van Figuur 3.5.
2. Waarom is de kans dat je Tom vindt met de "random friend of picked person" methode $\frac{5}{21}$. Geef een heel helder antwoord.
3. Programmeer de functie `kans_op_knoop(G,knoop)` die voor elk netwerk, voor elke knoop de kans berekent dat je die knoop trekt met de "random friend of picked person" methode.
4. Schrijf een test die checked dat die kansen optellen tot 1.
5. **Advanced** In de een na laatste paragraaf van 3.3 op blz 76 wordt een andere manier om die kansen te berekenen gemeld, die dezelfde uitkomst geeft. Programmeer die ook, en vergelijk de uitkomsten.
    

<!-- grade_id: FPa, points: 0, grade: false, solution: true -->
```python
Fig35 = nx.Graph()
Fig35.add_nodes_from(['Mary', 'Nancy', 'Tara', 'Bob', 'Pam', 'John', 'Tom'])
Fig35.add_edges_from([('Mary', 'Nancy'), ('Nancy', 'Bob'), ('Nancy', 'Tom'), ('Tara', 'Bob'), ('Bob', 'Tom'), ('Tom', 'Pam'), ('Tom', 'John'), ('Pam', 'John')])


def kans_op_knoop(G,knoop):
    totaal_knopen = len(G.nodes())
    degrees = sum([G.degree(x) for x in G.neighbors(knoop)])
    kans = (1 / totaal_knopen) * (1 / degrees)
    return kans

 

H= nx.karate_club_graph() 
# test
sum(kans_op_knoop(H,x) for x in H)
#kans_op_knoop(H,1)
```

**2** 
De kans dat een willekeurig knooppunt gekozen wordt bedraagt 1 op de 7. Nu willen we de kans berekenen om 'Tom' te selecteren als een willekeurige vriend van een willekeurig gekozen persoon. Tom heeft in totaal 4 vrienden (Bob, Pam, John en Nancy). De kans dat zijn vrienden gekozen worden en daarna Tom kiezen, bedraagt: 

Voor Bob is de kans (1/7) * (1/3) = 1/21.
Voor Pam is de kans (1/7) * (1/2) = 1/14.
Voor John is de kans (1/7) * (1/2) = 1/14.
Voor Nancy is de kans (1/7) * (1/3) = 1/21.

Dit bij elkaar opgeteld bedraagt: 5/21



<!-- grade_id: FPt, points: 1, grade: true, solution: false -->
```python
assert isinstance(Fig35, nx.Graph)
H= nx.karate_club_graph()
assert 0 <= kans_op_knoop(H,1)<= 1

```

# First course

Maak de sommen 3.11-3.24.

<!-- grade_id: fc, points: 14, grade: true, solution: true -->
```python
# Question 3.11
Fig310 = nx.Graph()
Fig310.add_nodes_from(['1', '2', '3', '4', '5', '6', '7'])
Fig310.add_edges_from([('1', '2'), ('1', '3'), ('2', '3'), ('3', '4'), ('2', '5'), ('5', '6'), ('5', '7')])
nx.draw(Fig310)
```

<!-- grade_id: fc324, points: 6, grade: true, solution: true -->
```python
#3.11
# Degree centrality
degree_centrality = nx.degree_centrality(Fig310)
highest_degree_node = max(degree_centrality, key=degree_centrality.get)
print("Knoop met de hoogste degree centrality:", highest_degree_node)

# Betweenness centrality
betweenness_centrality = nx.betweenness_centrality(Fig310)
highest_betweenness_node = max(betweenness_centrality, key=betweenness_centrality.get)
print("Knoop met de hoogste betweenness centrality:", highest_betweenness_node)

# Closeness centrality
closeness_centrality = nx.closeness_centrality(Fig310)
highest_closeness_node = max(closeness_centrality, key=closeness_centrality.get)
print("Knoop met de hoogste closeness centrality:", highest_closeness_node)
```

**3.12** -

**3.13** 
1. Shoe size of UK adults:

I would not expect to see a heavy-tailed distribution for shoe size of UK adults.Shoe sizes typically follow a more normal distribution where most people fall within a certain range of sizes, and there are fewer individuals with extremely small or large sizes.

2. US household income:

I would expect a heavy-tailed distribution. In many countries, including the US, a small number of high-income individuals or households account for a significant portion of the total income. This can be due to factors such as income inequality. 

3. Node degree in Twitter network:

In social networks like Twitter, I would expect to see a heavy-tailed distribution for the node degree. This is because popular or influential people attract more connections or followers. As a result, a few nodes (users) will have a significantly higher number of connections (followers), while most nodes will have only a few connections. This leads to a heavy-tailed distribution of node degrees.

4. Pairwise distance in the Wikipedia network:

I do not expect a heavy-tailed distribution. Most articles are reachable within a relatively small number of steps, with a diminishing number of articles requiring longer paths. Therefore, a heavy-tailed distribution is not expected for pairwise distances in the Wikipedia network.

**3.14** It would still be extremely rare to see a person with this height. While a heavy-tailed distribution allows for the possibility of extreme values, such as very tall individuals, the probability of encountering someone significantly outside the typical height range would still be extremely low.

**3.15** t/m **3.20** -

**3.21** I suspect that networks with heavy-tailed distributions are more vulnerable to targeted attacks. This can be due to the presence of highly connected hubs. Grid-like networks, on the other hand, are more vulnerable to random attacks because their regular structure makes them susceptible to disruptions that random removals can cause.

**3.22** D

**3.23** In terms of disrupting the network's connectivity and functionality, the node with the low clustering coefficient would be a better target. Its removal would result in a more significant disruption and potentially lead to the fragmentation of the network into smaller disconnected components.Removing or disabling this node would disconnect a relatively large portion of the network, causing fragmentation or isolation of different components.


```python

```