# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__ : D

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 1

# Deel 1: Programmeren

Voor deze opdrachten maken we gebruik van Networkx, de documentatie daarvan is te vinden via [deze link](https://networkx.github.io/documentation/stable/). Kijk vooral naar welke functies zij bieden, zodat je niet dingen doet die zij al voor je doen (tenzij dat natuurlijk de opdracht is).

Staat er `ModuleNotFoundError: No module named X`, doe dan in een cell `!pip install X` (natuurlijk is X hier iets van `matplotlib`, `networkx`, `scipy` of `nose`) en run de cell dan opnieuw.

Mocht dat niet werken, probeer dan `!pip3 install X` (werkt alleen bij Mac) of restart de kernel en probeer opnieuw! Mocht dat allemaal ook niet werken, roep dan even een TA!

```python
!pip install nose
```

```python
# Installeer benodigde packages
%matplotlib inline
import networkx as nx
import matplotlib.pyplot as plt
import numpy as np
import scipy
from nose.tools import assert_equal
from numpy.testing import assert_almost_equal, assert_array_almost_equal
```

## Vraag 1
<p>Maak een networkx ongerichte graaf `G` met de volgende knopen en lijnen:</p>
<center>
    V = {1, 2, 3, 4, 5} en 
    E = { (1, 2) , (1, 4) , (2, 3) , (2, 4) , (2, 5) , (3, 4) , (3, 5) }
</center>


<!-- grade_id: Vraag1, points: 0, grade: false, solution: true -->
```python
G = None

#WRITE YOUR CODE HERE
G = nx.Graph()
nodes_to_add = [1, 2, 3, 4, 5]
G.add_nodes_from(nodes_to_add)
edges_to_add = [(1, 2), (1, 4), (2, 3), (2, 4), (2, 5), (3, 4), (3, 5)]
G.add_edges_from(edges_to_add)

nx.draw(G, with_labels=True)
```

<!-- grade_id: Vraag1_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(nx.is_directed(G), False)
```

## Vraag 2
Maak op basis van het netwerk `G` een dictionary genaamd `neighborDict` met als key een node en als value de verzameling  met alle buren van die knoop.
Gebruik dict comprehensie.

Bijvoorbeeld: `neighborDict[5] = {2, 3}`

<!-- grade_id: Vraag2, points: 0, grade: false, solution: true -->
```python
neighborDict = None
#WRITE YOUR CODE HERE
neighborDict = {node:set(G.neighbors(node)) for node in G.nodes}
```

<!-- grade_id: Vraag2_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(neighborDict), dict)

```

## Vraag 3
Maak een functie `make_complete(G)` die teruggeeft hoeveel lijnen toegevoegd moeten worden aan een willekeurig  ongericht netwerk `G` (die misschien al lijnen bevat) om er een volledige graaf van te maken. Zorg dat het resultaat een integer is.


<!-- grade_id: Vraag3, points: 0, grade: false, solution: true -->
```python
def make_complete(G):
    #WRITE YOUR CODE HERE
    
    number_of_degrees_per_node = G.number_of_nodes() - 1
    total_degrees_needed = G.number_of_nodes() * number_of_degrees_per_node
    
    total_current_degrees = 0
    for node in G.nodes:
        total_current_degrees += G.degree(node)
    
    degrees_to_add = total_degrees_needed - total_current_degrees
    return degrees_to_add
    
# test 1
print(make_complete(G))
# test 2
ER = nx.gnm_random_graph(100, 200) # Zie https://networkx.github.io/documentation/stable/auto_examples/graph/plot_erdos_renyi.html
nx.draw(ER, with_labels=True);
make_complete(ER)
```

<!-- grade_id: Vraag3_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(make_complete(G)), int)
```

## Vraag 4
Maak een functie `my_degree(G)` die een dictionary maakt met als sleutels alle knopen van `G`    en als value de degree van de sleutel in `G`. Het resultaat is hetzelfde als `dict(G.degree())`.

* Gebruik een dict comprehensie!
* Je mag alleen de networkx functie `G.has_edge()` gebruiken. 

Zie: https://en.wikipedia.org/wiki/Degree_(graph_theory)

Voorbeeld resultaat op ons eigen netwerk `G`: `{1: 2, 2: 4...}`

<!-- grade_id: Vraag4, points: 0, grade: false, solution: true -->
```python
def my_degree(G):
    #WRITE YOUR CODE HERE
    
    
    return {k: sum(G.has_edge(k,n) for n in G) for k in G}
    
        
        
        
# test 1
print('G:', my_degree(G))

# # test 2
ER = nx.gnm_random_graph(10, 20) # Zie https://networkx.github.io/documentation/stable/auto_examples/graph/plot_erdos_renyi.html
nx.draw(ER, with_labels=True);
my_degree(ER)
```

<!-- grade_id: Vraag4_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(my_degree(G), dict(G.degree()))
```

## Vraag 5
Teken de subgraven van `G` geinduceerd door {1,2,3} en {2,5,3} en {5,4} en geef deze graven de namen respectievelijk `I`, `J` en `K`

<!-- grade_id: Vraag5, points: 0, grade: false, solution: true -->
```python

#WRITE YOUR CODE HERE
I = G.subgraph([1,2,3])
J = G.subgraph([2,5,3])
K = G.subgraph([5,4])

plt.figure(1)
nx.draw(I)
plt.figure(2)
nx.draw(J)
plt.figure(3)
nx.draw(K)
plt.show()
```

<!-- grade_id: Vraag5_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(sorted(list(I.nodes())), [1, 2, 3])
assert_equal(sorted(list(J.nodes())), [2, 3, 5])
assert_equal(sorted(list(K.nodes())), [4, 5])
```

## Vraag 6
Maak de functie `adj_matrix(G)`  en `adj_matrix2(G)` die allebei een adjacency matrix maken. In het geval van de graaf `G` die je hebt gemaakt bij vraag 1, heeft de matrix van deze  vorm (alleen dan in de vorm van een lijst van lijsten):
    <pre>
        1 2 3 4 5
      1 0 1 0 1 0
      2
      3
      4
      5
    </pre>
    
Een aantal eisen:
* Zorg dat de functies werken voor elk netwerk, niet alleen `G`. (We kijken de functies ook na met andere netwerken dan `G`.
* Gebruik enkel list comprehensies.
* Je mag in beide functies niet de networkx functie `nx.adjacency_matrix` gebruiken.
* Bij `adj_matrix(G)` moet je de networkx functie `G.has_edge` gebruiken en  geen enkele andere networkx functie.
* Bij `adj_matrix2(G)` moet je de networkx functie `G.neighbors` (dit is niet hetzelfde als `nx.neighbors`!) gebruiken en  geen enkele andere networkx functie.
* Zorg dat de nodes altijd gesorteerd zijn, zodat de matrix op volgorde is
* Hint:`int(True)==1` en `int(False)==0`, en dus ook `int(5==4)==0`.
    
Je antwoord voor netwerk `G` zal er zo uit zien:

```
[[0, 1, 0, 1, 0],
       ....
       ] 
```

<!-- grade_id: Vraag6, points: 0, grade: false, solution: true -->
```python
def adj_matrix(G):
    #WRITE YOUR CODE HERE
    return [[int(G.has_edge(i, j)) for j in G] for i in G]
    

def adj_matrix2(G):
    #WRITE YOUR CODE HERE
    return [[int(j in list(G.neighbors(i))) for j in G] for i in G]
    

adj_matrix(G), adj_matrix2(G), adj_matrix(G)==adj_matrix2(G)

# test het zelf ook met de ER graaf.
```

<!-- grade_id: Vraag6_test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(adj_matrix(G)), list)
assert_equal(type(adj_matrix(G)[0]), list)
assert_equal(type(adj_matrix(G)[0][0]), int)
assert_equal(type(adj_matrix2(G)), list)
assert_equal(type(adj_matrix2(G)[0]), list)
assert_equal(type(adj_matrix2(G)[0][0]), int)
assert_equal(type(adj_matrix2(G)[0][0]), int)
assert_array_almost_equal(adj_matrix(G), adj_matrix2(G))

```

# Opgaven uit First Course

1. Alles uit Hfd 0
2. Hfd 1: alles tot en met 21.
3. Maak opgave 21 in `networkx`
        * maak het netwerk, met `G.add_edges`
        * print het
        * maak ook de adjecency matrix als een lijst van lijsten
        * check dat deze adjencency matrix hetzelfde is als een adjececy matrix die networkx maakt.
4. Schrijf nu ook checks voor alle vragen 1.17-1.20 in Python en controleer je antwoord.

<!-- grade_id: FCa, points: 1, grade: true, solution: true -->
#Hfd 0

0.1
D

0.2
B

0.3
A link represents a flight between these two airports.

0.4
B

0.5
Undirected

0.6
Directed


#Hfd1

1.1 ✅

1.2:

- 2

- Lmax = N(N−1)/2 dus in dit geval N = 1 geeft Lmax = 0

1.3
Degree van 2

1.4
Beide hebben een degree van 1.

1.5
Size of the nodes

1.6
D

1.7
D

1.8
B

1.9
A, undirected cause dialogue is at least two people talking so it automatically goes both ways and unweighted cause the network represents not how much the people talk in the dialogue, only if they engage in a dialogue.

1.10
D

1.11
Complete network, (25* 24)/2 = 300 links

1.12
max links = N(N-1)/2

1.13
max links = N1 * N2

1.14
A

1.15
Yes because then each node is linked to every other node since N(N-1). 

1.16
D

1.17
6 nodes and 10 links, there are no self loops

1.18
There is a node with outgoing links to every other node = F
There are no nodes with ingoing links to every other node.

1.19
C

1.20
in -strength = 5
out-strength = 0

1.21
zie onderstaand bij #1.21


```python
#1.17
# directed weighted graph maken
DG = nx.DiGraph()

nodes = ['A', 'B', 'C', 'D', 'E', 'F']
DG.add_nodes_from(nodes)
DG.add_edge('A', 'B', weight = 1)
DG.add_edge('B', 'C', weight = 2)
DG.add_edge('D', 'B', weight = 1)
DG.add_edge('D', 'E', weight = 1)
DG.add_edge('E', 'F', weight = 1)
DG.add_edge('F', 'A', weight = 2)
DG.add_edge('F', 'B', weight = 1)
DG.add_edge('F', 'C', weight = 3)
DG.add_edge('F', 'D', weight = 1)
DG.add_edge('F', 'E', weight = 1)


# print het:

# om de positie van de nodes te berekenen, maps elke node met een x,y positie
pos = nx.spring_layout(DG)
# tekent nodes met behulp van hun positie
nx.draw_networkx_nodes(DG, pos)
# tekent lijnen tussen nodes met behulp van hun positie
nx.draw_networkx_edges(DG, pos)
# pakt weights van edges(node-node)
edge_labels = nx.get_edge_attributes(DG, "weight")
# geeft labels aan edges met behulp van edge_labels
nx.draw_networkx_edge_labels(DG, pos, edge_labels=edge_labels)
#labels van nodes laten zien
node_labels = {n: str(n) for n in DG}
nx.draw_networkx_labels(DG, pos, labels=node_labels)




```

```python
#checks voor 1.17

def number_of_nodes(DG):
    return len([node for node in DG])

def number_of_links(DG):
    links = 0
    for node in DG:
        links += DG.degree(node)
    return links/2

def self_loops(DG):
    for node in DG:
        if node in DG.neighbors(node):
            return node
        else:
            print('No self loop') 

```

```python
#checks voor 1.18

def node_outgoing_to_all_nodes(DG):
    for node in DG:
        if len(list(DG.neighbors(node))) == len(DG.nodes) - 1:
            return node

def node_ingoing_to_all_nodes(DG):
    for node in DG:
        if len(list(DG.predecessors(node))) == len(DG.nodes) - 1:
            return node
        else:
            print('No node')
```

```python
#checks voor 1.19
def node_is_sink(DG):
    for node in DG:
        if len(list(DG.predecessors(node))) > 0 and len(list(DG.neighbors(node))) == 0:
            return node
        else:
            continue
```

```python
#checks voor 1.20
def in_strength(DG):
    for node in DG:
        if node == 'C':
            return DG.degree(node, weight='weight')
            
def out_strength(DG):
    for node in DG:
        if node == 'C':
            return len(list(DG.successors(node)))
```

```python
#1.21
# maak de graph unweighted en undirected:
undirected_DG = DG.to_undirected()

nx.draw(undirected_DG, with_labels = True)


#nodes in this network = 6
#links in this network = 9
```

# FC vraag 1.13

* Maak deze vraag ook in networkx.
    * Maak een functie `Q113(N1,N2)` die een network teruggeeft zoals gevraagd
    * Met welke check check je dat het "volledig" (voor een bipartiet netwerk natuurlijk) is?
    
* **Hint** Google eens wat.

<!-- grade_id: fc113a, points: 0, grade: false, solution: true -->
```python
def Q113(N1,N2):
    '''Return a maximal bipartite network with N1 nodes on one side and N2 on the other.'''
    #WRITE YOUR CODE HERE
    BG = nx.Graph()
    BG.add_nodes_from(N1, bipartite = 0)
    BG.add_nodes_from(N2, bipartite = 1)
    BG.add_edges_from[()]
    
    l, r = nx.bipartite.sets(BG)
    
    bipartite = nx.draw(BG, with_labels = True)
    return bipartite
    
def maximaal_bipartiet(BG):
    '''Return True if BG is a maximal bipartite network. False in all other cases.'''
    #WRITE YOUR CODE HERE
    for node in BG.nodes():
        
        BG.neighbors()
    
    
    
#test
maximaal_bipartiet(Q113(5,10))
```

<!-- grade_id: fc113t, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q113(5,10),nx.Graph)
assert isinstance(maximaal_bipartiet(Q113(5,10)), bool)


    
 
```

<!-- grade_id: adjmatrixchecksd, points: 1, grade: true, solution: true -->
```python
#WRITE YOUR CODE HERE
```

<!-- grade_id: FCa21, points: 1, grade: true, solution: true -->
```python
#WRITE YOUR CODE HERE
```

# Verzamelingen

* Dubbel klik eens op deze cel zodat je de code ziet.
* Leer dan de latex codes voor de mooie symbolen, en zorg dat je zelf nu ook netjes antwoord kunt geven.
* Formules zet je altijd tussen dollartekens.
    * Als je het in het midden van de regel wil tussen dubbele dolalartekens
* Elk commando begint altijd met `\`. Spaties maken niks uit. Latex regelt hoe het er netjes uit komt te zien.

### Opgaven

1. $\{1,2,3\}\cap\{2,3,4\}=?$
1. $\{1,2,3\}\cup\{2,3,4\}=?$
1. $\{1,2,3\}\setminus\{2,3,4\}=?$
1. $\emptyset\setminus\{2,3,4\}=?$
2. $\{1,2,3\}\times\{2,3,4\}=?$
3. De machtsverzameling van $\{1,2,3,4\}$ is ....
     * Check dat je echt $2^4$ veel deelverzamelingen hebt gevonden.
     * het heet niet voor niks *machts* verzameling ;-)

#### Teken ook netjes de Venn diagrammen
#### Check je antwoorden ook in Python
#### Maak zelf opgaven en leg die aan je groepsleden voor

# Comprehensies

* Zorg dat je list, set en dict comprehensies echt goed in je vingers krijgt!
* Het is heerlijk als je het kunt.

### Goede uitleg

Ik [Googlede op python comprehension tutorial](https://www.google.com/search?q=python+comprehension+tutorial) en de [video](https://www.youtube.com/watch?v=3dt4OGnU5sM) is prima, maar ook de [datacamp tutorial](https://www.datacamp.com/community/tutorials/python-list-comprehension) of de [real python](https://realpython.com/list-comprehension-python/) tutorial.

### Wat opgaven

Je gaat allemaal comprehensies doen op een gedicht van annie MG Schmidt. Het staat in de cel hieronder als string in de variabele `AMGS` en daarna als lijst van woorden in `t`.

In je tests gebruik je allemaal string functies, en ook `in`. 

#### list comprehensies

1. de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)

2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)

3. de lijst van alle woorden in `AMGS` van lengte meer dan 3
4. de lijst van alle woorden in `AMGS` waar een `a` in zit
5. nu waar juist geen `a` in zit
6. nu waar als er een `a` in zit ook een `b` in zit (wat betekent dit precies he?)
7. alle woorden zonder klinker.
8. Voor elk woord zijn lengte.
9. De lijst met het hele gedicht maar nu waar elk woord met een hoofdletter begint.
10. De lijst die net zo lang is als het gedicht (dus `t`), maar dan op elke positie of het woord op die positie gelijk is aan `ik`. Dus de waardes zijn alleen maar `True` en `False`.
    * Gebruik `sum` om te berekenen hoe vaak het woord `ik` voorkomt.
11. Verzien er zelf 10 bij, en maak die ook, en geef ze aan je groepsgenoeten.

#### set comprehensies

1. De verzameling van alle unieke woorden in `AMGS`.
2. De verzameling van alle unieke woorden in `AMGS` met `ch` er in.
3. De verzameling van alle unieke woorden in `AMGS`, maar dan geschreven in hoofdletters.
4. Alle woordlengtes die in `AMGS` voorkomen.


### Dict comprehensies

1. De dict met voor elk uniek "woord" in `AMGS` zijn lengte. Dus het woord is de key en de lengte de waarde. Let op , woorden zijn dus de elementen in `t`. (antwoord beneden)
2. De dict met voor elk uniek "woord" in `AMGS` de verzameling karakters in dat woord.
3. De dict met voor elk uniek "woord" in `AMGS`  de "schoongemaakte" versie van dat woord (zonder leestekens, etc, gebruik een lange lijst als `a.replace(':','').replace(',','')`. Zet het woord ook in klein eletters.

    
### Nesting

3. De dict met voor elk uniek woord, de verzameling klinkers in dat  woord. 
    *   (antwoord beneden, probeer eerst minstens een half uur zelf)
4.  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.
    (antwoord ook beneden, niet meteen gan kijken.)
5. DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!   Hint loop over de posities. 
6. Na de vorige is deze heel makkelijk: DE dict met voor elk woord  hoe vaak het voorkomt. 


```python
#1. de lijst van alle woorden in `AMGS` beginnend met een `i`.
mijnlisti = [x for x in t if x.startswith("i")]
```

```python
#2. zelfde maar dan beginnend met `i` of `I`.
1 
mijnlistiI_1 = [x for x in t if x.startswith("i") or x.startswith("I")]
2
mijnlistiI_2 = [word for word in t if word.lower().startswith('i')]
3
mijnlistiI_3 = [w if w.startswith('i') else w if w.startswith('I') else '' for w in t]
mijnlistiI_3_filtered = list(filter(lambda x: x != '', words_i))
```

```python
#3. de lijst van alle woorden in `AMGS` van lengte meer dan 3
mijnlistlen = [x for x in t if (len(x)>2)]
```

```python
#4. de lijst van alle woorden in `AMGS` waar een `a` in zit
mijnlista= [x for x in t if "a" in x]
```

```python
#5. nu waar juist geen `a` in zit
mijnlistgeena = [x for x in t if "a" not in x]
```

```python
#6. nu waar als er een `a` in zit ook een `b` in zit (wat betekent dit precies he?)
[w if 'a' and 'b' in t else w for w in t]
```

```python
#7. alle woorden zonder klinker.
noklink = [x for x in t if 
           "a" not in x and
           "i" not in x and
           "I" not in x and
           "A" not in x and
           "E" not in x and
           "e" not in x and
           "O" not in x and 
           "o" not in x and
           "u" not in x and
           "U" not in x]
```

```python
#8. Voor elk woord zijn lengte.
lengtelist = [len(x) for x in t]
```

```python
#9. De lijst met het hele gedicht maar nu waar elk woord met een hoofdletter begint.
#hele gedicht waar elk eerste letter een hoofdletter is.
lijsthoofd = [x.capitalize() for x in t]
```

```python
# 10. De lijst die net zo lang is als het gedicht (dus `t`), maar dan op elke positie of het woord op die positie gelijk is aan `ik`. Dus de waardes zijn alleen maar `True` en `False`.
lijstik = [True if x == "ik" else False for x in t]

#aantal keer "ik/True"
print(sum(True if x == "ik" else False for x in t))

```

```python
# 11. Verzien er zelf 10 bij, en maak die ook, en geef ze aan je groepsgenoeten.
#11.1 geef een lijst waar het hele gedicht in staat zonder hoofdletters
lijst111 = x,[x.lower() for x in t]
lijst111
#11.2 geef een lijst waar alle woorden in staan die met al beginnen.
lijst112 = [x for x in t if x.startswith("al")]
lijst112
#11.3 geef een lijst waar alle woorden in staan die eindigen op een l.
lijst113 = [x for x in t if x.endswith("l")]
lijst113
#11.4 geef een lijst waar "ik" wordt vervangen door "jan"
lijst114 = ["jan" if x == "ik" else x for x in t]
lijst114
#11.5 maak een list waar elk woord maar 1x voorkomt.
list115 = list({x for x in t})
list115
#11.6 maak een list wmet elk woord dat bestaat uit een even aantal letters en begint met een a.
list116 =  [x for x in t if len(x) % 2 == 0 and x.startswith("a")]
list116
#11.7 maak een list met elk woord dat bestaat uit een even aantal letters.
list117 =  [x for x in t if len(x) % 2 == 0]
list117
#11.8 maak een lijst waar elke instantie van het woord "niet" wordt vervangen met de smiley ":)""
list118 =  [":)" if x == "niet" else x for x in t]
list118
#11.9 maak een lijst waar het woord bestaat uit een oneven aantal letters/leestekens
list119 =  [x for x in t if len(x) % 2 != 0]
list119
#11.10, maak een lijst die ":)" returnd als een woord een hoofdletter bevat, anders ":("
list1110 = [":)" if any(letter.isupper() for letter in x) else ":(" for x in t]
list1110
```

```python
# 1. De verzameling van alle unieke woorden in `AMGS`.
set1 = {x for x in t}
set1
```

```python
# 2. De verzameling van alle unieke woorden in `AMGS` met `ch` er in.
set2 = {x for x in t if "ch" in x}
set2
```

```python
# 3. De verzameling van alle unieke woorden in `AMGS`, maar dan geschreven in hoofdletters.
set3 = {x.upper() for x in t}
set3
```

```python
# 4. Alle woordlengtes die in `AMGS` voorkomen.
set4 = {len(x) for x in t}
set4
```

```python
# 1. De dict met voor elk uniek "woord" in `AMGS` zijn lengte. Dus het woord is de key en de lengte de waarde. Let op , woorden zijn dus de elementen in `t`. (antwoord beneden)
#woord is key, lengte is info, in een dict
dict1 = {x:{len(x)}for x in t}
dict1
```

```python
# 2. De dict met voor elk uniek "woord" in `AMGS` de verzameling karakters in dat woord.
#het woord als key, vervolgens het aantal unieke characters
dict2 = {x: set(x) for x in t}
dict2
```

```python
# 3. De dict met voor elk uniek "woord" in `AMGS`  de "schoongemaakte" versie van dat woord (zonder leestekens, etc, gebruik een lange lijst als `a.replace(':','').replace(',','')`. Zet het woord ook in klein eletters.
#een dict die de leestekens en hoofdletters verdwijdert
def schoonmaken(woord):
    leestekens = '!"#$%&\'()*+,-./:;<>?@[\\]^_`{|}~='
    return woord.translate(str.maketrans('', '', leestekens)).lower()

dict3 = {woord: schoonmaken(woord) for woord in t}
dict3
```

```python
# 3. De dict met voor elk uniek woord, de verzameling klinkers in dat  woord. 

dict1a = {x:{y for y in x if y in 'aeiouAEIOU'} for x in set(t)}
dict1a
```

```python
# 4.  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.

dict5 = {x:{t[x-1],t[x+1]} for x in range(1,len(t)-1)}
dict5
```

```python
# 5. DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!   Hint loop over de posities.
#key is het woord, vervolgens de posities van het woord in het gedicht
dict6 = {x: {y for y, z in enumerate(t) if z == x} for x in t}
dict6
```

```python
# 6. Na de vorige is deze heel makkelijk: DE dict met voor elk woord  hoe vaak het voorkomt. 
#key is het woord, vervolgens hoe vaak het woord voorkomt
dict7 = {x: t.count(x) for x in t}
dict7
```

```python
AMGS='''
Ik wil niet meer, ik wil niet meer!
Ik wil geen handjes geven!
Ik wil niet zeggen elke keer:
Jawel mevrouw, jawel meneer...
nee, nooit meer in m'n leven!
Ik hou m'n handen op m'n rug
en ik zeg lekker niks terug!

Ik wil geen vieze havermout,
ik wil geen tandjes poetsen!
Ik wil lekker knoeien met het zout,
ik wil niet aardig zijn, maar stout
en van de leuning roetsen
en schipbreuk spelen in de teil
en ik wil spugen op het zeil!

En heel hard stampen in een plas
en dan m'n tong uitsteken
en morsen op m'n nieuwe jas
en ik wil overmorgen pas
weer met twee woorden spreken!
En ik wil alles wat niet mag,
de hele dag, de hele dag!

En ik wil op de kanapee
met hele vuile schoenen
en ik wil aldoor gillen: nee!
En ik wil met de melkboer mee
en dan het paardje zoenen.
En dat is alles wat ik wil
en als ze kwaad zijn, zeg ik: Bil!'''

t=AMGS.split()
len(t),t[:10]
```

```python
# probeer eens alleen t[0]. en dan shift return, en probeer eens wat van de suggesties
# om erachter te komen wat iets doet typ je t[0].isupper? en dan shift enter.
# Ideaal he! Dat kan bij elk object. t[0] is natuurlijk een string, dus je krijgt string functies te zien. 
t[0].isupper()
```

```python
# "in" is ook heel handig in Python

print(t[1])
'i' in t[1], 'W' in t[1], 'W' in t[1].upper()
```

```python
#de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)
[w for w in t if w.startswith('i')]
```

```python
#2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)
# Verzin minstens 2 andere manieren om exact hetzelfde te krijgen
[w for w in t if w.startswith('i') or w.startswith('I')]
```

```python
#  De dict met voor elk uniek woord in `AMGS` zijn lengte.

{w:len(w) for w in t}
```

```python
# De dict met voor elk uniek woord, de verzameling klinkers in dat  woord.

{w:{l for l in w if l in 'aeiou'} 
    for w in set(t)} # verbeter hem zelf zodat hoofdletters ook mee worden genomen

```

```python
#  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.

{i:{t[i-1],t[i+1]} for i in range(1,len(t)-1)}
```

<!-- grade_id: TRICKY, points: 0, grade: false, solution: true -->
```python
#  DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!    
l=dict()
#WRITE YOUR CODE HERE
```

```python

```