# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__ : F

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 1

# Deel 1: Programmeren

Voor deze opdrachten maken we gebruik van Networkx, de documentatie daarvan is te vinden via [deze link](https://networkx.github.io/documentation/stable/). Kijk vooral naar welke functies zij bieden, zodat je niet dingen doet die zij al voor je doen (tenzij dat natuurlijk de opdracht is).

Staat er `ModuleNotFoundError: No module named X`, doe dan in een cell `!pip install X` (natuurlijk is X hier iets van `matplotlib`, `networkx`, `scipy` of `nose`) en run de cell dan opnieuw.

Mocht dat niet werken, probeer dan `!pip3 install X` (werkt alleen bij Mac) of restart de kernel en probeer opnieuw! Mocht dat allemaal ook niet werken, roep dan even een TA!

```python
# Installeer benodigde packages
%matplotlib inline
import networkx as nx
import matplotlib.pyplot as plt
import numpy as np
import scipy
from nose.tools import assert_equal
from numpy.testing import assert_almost_equal, assert_array_almost_equal
```

## Vraag 1
<p>Maak een networkx ongerichte graaf `G` met de volgende knopen en lijnen:</p>
<center>
    V = {1, 2, 3, 4, 5} en 
    E = { (1, 2) , (1, 4) , (2, 3) , (2, 4) , (2, 5) , (3, 4) , (3, 5) }
</center>


<!-- grade_id: Vraag1, points: 0, grade: false, solution: true -->
```python
G = None


#WRITE YOUR CODE HERE
G = nx.Graph([(1, 2) , (1, 4) , (2, 3) , (2, 4) , (2, 5) , (3, 4) , (3, 5)])
nx.draw_networkx(G, with_labels=True)
```

<!-- grade_id: Vraag1_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(nx.is_directed(G), False)
```

## Vraag 2
Maak op basis van het netwerk `G` een dictionary genaamd `neighborDict` met als key een node en als value de verzameling  met alle buren van die knoop.
Gebruik dict comprehensie.

Bijvoorbeeld: `neighborDict[5] = {2, 3}`

<!-- grade_id: Vraag2, points: 0, grade: false, solution: true -->
```python
neighborDict = None
#WRITE YOUR CODE HERE
neighborDict = {1:(2,4), 2:(1,3,4,5), 3:(2,4,5), 4:(1,2,3), 5:(2,3)}
```

<!-- grade_id: Vraag2_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(neighborDict), dict)

```

## Vraag 3
Maak een functie `make_complete(G)` die teruggeeft hoeveel lijnen toegevoegd moeten worden aan een willekeurig  ongericht netwerk `G` (die misschien al lijnen bevat) om er een volledige graaf van te maken. Zorg dat het resultaat een integer is.


<!-- grade_id: Vraag3, points: 0, grade: false, solution: true -->
```python
def make_complete(G):
    "Geeft terug hoeveel lijnen worden toegevoegd om een volledige graaf te maken"
    total_added_lines = 0
    for node in G:
        relation_count = 0 #count om het aantal relaties van de node te tellen
        for edge in G.edges():
            if edge[0] == node or edge[1] == node:
                #telt of de edge een link met de node bevat
                relation_count += 1
        total_added_lines += (len(G.nodes())-1) - relation_count #voegt aantal nodige relaties bij het totaal op
    return total_added_lines
            
    

# test 1
print(make_complete(G))
# test 2
ER = nx.gnm_random_graph(100, 200) # Zie https://networkx.github.io/documentation/stable/auto_examples/graph/plot_erdos_renyi.html
nx.draw_networkx(ER, with_labels=True);
make_complete(ER)
```

<!-- grade_id: Vraag3_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(make_complete(G)), int)
```

## Vraag 4
Maak een functie `my_degree(G)` die een dictionary maakt met als sleutels alle knopen van `G`    en als value de degree van de sleutel in `G`. Het resultaat is hetzelfde als `dict(G.degree())`.

* Gebruik een dict comprehensie!
* Je mag alleen de networkx functie `G.has_edge()` gebruiken. 

Zie: https://en.wikipedia.org/wiki/Degree_(graph_theory)

Voorbeeld resultaat op ons eigen netwerk `G`: `{1: 2, 2: 4...}`

<!-- grade_id: Vraag4, points: 0, grade: false, solution: true -->
```python
def my_degree(G):
    "returns dict with {node:degree} pairs for graph G"
    return {node1: sum([int(G.has_edge(*(node1, node2))) for node2 in G]) for node1 in G}

# test 1
print('G:', my_degree(G))

# test 2
ER = nx.gnm_random_graph(10, 20) # Zie https://networkx.github.io/documentation/stable/auto_examples/graph/plot_erdos_renyi.html
nx.draw_networkx(ER, with_labels=True);
my_degree(ER)
```

<!-- grade_id: Vraag4_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(my_degree(G), dict(G.degree()))
```

## Vraag 5
Teken de subgraven van `G` geinduceerd door {1,2,3} en {2,5,3} en {5,4} en geef deze graven de namen respectievelijk `I`, `J` en `K`

<!-- grade_id: Vraag5, points: 0, grade: false, solution: true -->
```python
I = None
J = None
K = None

#WRITE YOUR CODE HERE
I = G.subgraph([1,2,3])
J = G.subgraph([5,2,3])
K = G.subgraph([5,4])

plt.figure(1)
nx.draw_networkx(I)
plt.figure(2)
nx.draw_networkx(J)
plt.figure(3)
nx.draw_networkx(K)
plt.show()
```

<!-- grade_id: Vraag5_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(sorted(list(I.nodes())), [1, 2, 3])
assert_equal(sorted(list(J.nodes())), [2, 3, 5])
assert_equal(sorted(list(K.nodes())), [4, 5])
```

## Vraag 6
Maak de functie `adj_matrix(G)`  en `adj_matrix2(G)` die allebei een adjacency matrix maken. In het geval van de graaf `G` die je hebt gemaakt bij vraag 1, heeft de matrix van deze  vorm (alleen dan in de vorm van een lijst van lijsten):
    <pre>
        1 2 3 4 5
      1 0 1 0 1 0
      2
      3
      4
      5
    </pre>
    
Een aantal eisen:
* Zorg dat de functies werken voor elk netwerk, niet alleen `G`. (We kijken de functies ook na met andere netwerken dan `G`.
* Gebruik enkel list comprehensies.
* Je mag in beide functies niet de networkx functie `nx.adjacency_matrix` gebruiken.
* Bij `adj_matrix(G)` moet je de networkx functie `G.has_edge` gebruiken en  geen enkele andere networkx functie.
* Bij `adj_matrix2(G)` moet je de networkx functie `G.neighbors` (dit is niet hetzelfde als `nx.neighbors`!) gebruiken en  geen enkele andere networkx functie.
* Zorg dat de nodes altijd gesorteerd zijn, zodat de matrix op volgorde is
* Hint:`int(True)==1` en `int(False)==0`, en dus ook `int(5==4)==0`.
    
Je antwoord voor netwerk `G` zal er zo uit zien:

```
[[0, 1, 0, 1, 0],
       ....
       ] 
```

<!-- grade_id: Vraag6, points: 0, grade: false, solution: true -->
```python
def adj_matrix(G):
    ""
    sorted_nodes = sorted([node for node in G])
    return [[1 if G.has_edge(*(node1, node2)) else 0 for node2 in sorted_nodes]for node1 in sorted_nodes]

def adj_matrix2(G):
    ""
    sorted_nodes = sorted([node for node in G])
    return [[int(node2 in G.neighbors(node1)) for node2 in sorted_nodes] for node1 in sorted_nodes]

adj_matrix(G), adj_matrix2(G), adj_matrix(G)==adj_matrix2(G)

# test het zelf ook met de ER graaf.
HR = nx.gnm_random_graph(7, 16) 
nx.draw_networkx(HR)
adj_matrix(HR), adj_matrix2(HR), adj_matrix(HR)==adj_matrix2(HR)
```

<!-- grade_id: Vraag6_test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(adj_matrix(G)), list)
assert_equal(type(adj_matrix(G)[0]), list)
assert_equal(type(adj_matrix(G)[0][0]), int)
assert_equal(type(adj_matrix2(G)), list)
assert_equal(type(adj_matrix2(G)[0]), list)
assert_equal(type(adj_matrix2(G)[0][0]), int)
assert_equal(type(adj_matrix2(G)[0][0]), int)
assert_array_almost_equal(adj_matrix(G), adj_matrix2(G))

```

# Opgaven uit First Course

1. Alles uit Hfd 0
2. Hfd 1: alles tot en met 21.
3. Maak opgave 21 in `networkx`
        * maak het netwerk, met `G.add_edges`
        * print het
        * maak ook de adjecency matrix als een lijst van lijsten
        * check dat deze adjencency matrix hetzelfde is als een adjececy matrix die networkx maakt.
4. Schrijf nu ook checks voor alle vragen 1.17-1.20 in Python en controleer je antwoord.

<!-- grade_id: FCa, points: 1, grade: true, solution: true -->
Hoofdstuk 0:<br>
0.1: D<br>
0.2: B<br>
0.3: A direct flight between these two airports.<br>
0.4: D <br>
0.5: Undirected <br>
0.6: Directed <br><br>
Hoofdstuk 1:<br>
1.1: Door Chapter 1 tutorial van github gegaan<br>
1.2: Eén link kan maximaal twee knopen connecten, Eén knoop kan een maximaal (N-1) links hebben naar andere knopen.<br>
1.3: Een degree van 4<br>
1.4: Meest gebruikelijke In-degree is 2 en de meest gebruikelijke out-degree is 2<br>
1.5: Het gewicht van een link kan hiervoor gebruikt worden<br>
1.6: C<br>
1.7: D<br>
1.8: B<br>
1.9: A, want het gaat niet om hoeveel dialoog ze hebben, maar gewoonweg of ze überhaupt een dialoog voeren. Het is dus unweighted<br>
1.10: D<br>
1.11: Een clique<br>
1.12: maximum: N(N-1)/2<br>
1.13: maximum: N1 * N2<br>
1.14: A<br>
1.15: Ja deze is altijd compleet. Tenzij het netwerk meerdere links tussen dezelfe knopen toelaat, of een knoop een link met zichzelf kan hebben<br>
1.16: D<br>
1.17: 6 knopen, 10 links, geen self-loops<br>
1.18: Ja, knoop F heeft uitgaande links met alle andere knopen. Geen enkele knoop heeft ingaande links van alle knopen<br>
1.19: Knoop C is een sink-link<br>
1.20: In-strength: 5, out-strength: 0<br>
1.21: 6 knopen en 9 links

```python
#Voor opgaven 1.17 tot 1.20
G17_20 = nx.DiGraph()
G17_20.add_weighted_edges_from([('A','B',1), ('B','C',2), ('D','B',1), ('D','E',1), ('E','F',1), ('F','A',2), ('F','B',1), ('F','C',3), ('F','D',1), ('F','E',1)])
nx.draw_networkx(G17_20)

#Antwoorden checken van 1.17 tot 1.20
#1.17
assert len(G17_20.nodes()) == 6
assert len(G17_20.edges()) == 10
assert len([node for node in G17_20 if node in list(G17_20.neighbors(node)) or node in list(G17_20.predecessors(node))]) == 0
#1.18
assert len(list(G17_20.neighbors('F'))) == len(list(G17_20.nodes()))-1
#1.19
assert len(list(G17_20.neighbors('C'))) == 0 and len(list(G17_20.predecessors('C'))) >= 1
#1.20
assert G17_20.in_degree('C', weight='weight') == 5 and G17_20.out_degree('C', weight='weight') == 0
```

```python
#Voor vraag 1.21
G21 = nx.Graph()
G21.add_edges_from([("A","B"), ("A","F"), ("B","C"), ("B","D"), ("B","F"), ("C","F"), ("D","E"), ("D","F"), ("E","F")])
adj_matrix_G21 = [[0, 1, 0, 0, 0, 1], [1, 0, 1, 1, 0, 1], [0, 1, 0, 0, 0, 1], [0, 1, 0, 0, 1, 1], [0, 0, 0, 1, 0, 1], [1, 1, 1, 1, 1, 0]]

#visualising
print(G21)
nx.draw_networkx(G21)
```

```python
#Kan niet vergelijken met nx.adjecency_matrix(G21), want dit levert een scipy.sparse array op in Networkx 3.0
#Wordt daarom vergeleken met eerder geschreven adj_matrix() en adj_matrix2() functie

assert adj_matrix_G21 == adj_matrix(G21) == adj_matrix2(G21)
```

# FC vraag 1.13

* Maak deze vraag ook in networkx.
    * Maak een functie `Q113(N1,N2)` die een network teruggeeft zoals gevraagd
    * Met welke check check je dat het "volledig" (voor een bipartiet netwerk natuurlijk) is?
    
* **Hint** Google eens wat.

<!-- grade_id: fc113a, points: 0, grade: false, solution: true -->
```python
def Q113(N1,N2):
    '''Return a maximal bipartite network with N1 nodes on one side and N2 on the other.'''
    return nx.Graph([edge for edgelist in [[(n1, n2) for n1 in range(N1)] for n2 in range(N1, N1+N2)] for edge in edgelist])

def maximaal_bipartiet(BG):
    '''Return True if BG is a maximal bipartite network. False in all other cases.'''
    neighbor_list = []
    for node in BG:
        if list(BG.neighbors(node)) not in neighbor_list:
            neighbor_list.append(list(BG.neighbors(node)))
    return len(neighbor_list) == 2


#test
nx.draw_networkx(Q113(5,10))
maximaal_bipartiet(Q113(5,10))
```

<!-- grade_id: fc113t, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q113(5,10),nx.Graph)
assert isinstance(maximaal_bipartiet(Q113(5,10)), bool)


    
 
```

<!-- grade_id: adjmatrixchecksd, points: 1, grade: true, solution: true -->
```python

```

<!-- grade_id: FCa21, points: 1, grade: true, solution: true -->
```python

```

# Verzamelingen

* Dubbel klik eens op deze cel zodat je de code ziet.
* Leer dan de latex codes voor de mooie symbolen, en zorg dat je zelf nu ook netjes antwoord kunt geven.
* Formules zet je altijd tussen dollartekens.
    * Als je het in het midden van de regel wil tussen dubbele dolalartekens
* Elk commando begint altijd met `\`. Spaties maken niks uit. Latex regelt hoe het er netjes uit komt te zien.

### Opgaven

1. $\{1,2,3\}\cap\{2,3,4\}=\{2,3\}$
1. $\{1,2,3\}\cup\{2,3,4\}=\{1,2,3,4\}$
1. $\{1,2,3\}\setminus\{2,3,4\}=\{1\}$
1. $\emptyset\setminus\{2,3,4\}=\emptyset$
2. $\{1,2,3\}\times\{2,3,4\}=\{(1,2), (1,3), (1,4), (2,2), (2,3), (2,4), (3,2), (3,3), (3,4)\}$
3. De machtsverzameling van $\{1,2,3,4\}$ is $\{\{1\}, \{2\}, \{3\}, \{4\}, \{1,2\}, \{1,3\}, \{1,4\}, \{2,3\}, \{2,4\}, \{3,4\}, \{1,2,3\}, \{1,2,4\}, \{1,3,4\}, \{2,3,4\}, \{1,2,3,4\}, \emptyset\} $
     * Check dat je echt $2^4$ veel deelverzamelingen hebt gevonden.
     * het heet niet voor niks *machts* verzameling ;-)

#### Teken ook netjes de Venn diagrammen
#### Check je antwoorden ook in Python
#### Maak zelf opgaven en leg die aan je groepsleden voor

```python
set1 = set('123')
set2 = set('234')

#check voor vraag 1:
assert (set1 & set2) == set('23')
#check voor vraag 2:
assert (set1 | set2) == set('2341')
#check voor vraag 3:
assert (set1 - set2) == set('1')
#check voor vraag 4:
assert (set() - set2) == set()

#Dit is om vraag 5 & 6 te checken
from itertools import product, chain, combinations
#Deze functie is gehaald van de documentatiepagina van itertools om functionaliteit uit te breiden: https://docs.python.org/3/library/itertools.html
def powerset(iterable): 
    "powerset([1,2,3]) --> () (1,) (2,) (3,) (1,2) (1,3) (2,3) (1,2,3)"
    s = list(iterable)
    return chain.from_iterable(combinations(s, r) for r in range(len(s)+1))

#check voor vraag 5:
assert set(product(set1, set2)) == {('1','2'), ('1','3'), ('1','4'), ('2','2'), ('2','3'), ('2','4'), ('3','2'), ('3','3'), ('3','4')}
#check voor vraag 6:
assert set((powerset({1,2,3,4}))) == {(1, 3), (2, 4), (2,), (1, 2), (3, 4), (1, 2, 3), (4,), (1, 4), (2, 3), (1,), (1, 3, 4), (1, 2, 3, 4), (2, 3, 4), (), (3,), (1, 2, 4)}
```

# Comprehensies

* Zorg dat je list, set en dict comprehensies echt goed in je vingers krijgt!
* Het is heerlijk als je het kunt.

### Goede uitleg

Ik [Googlede op python comprehension tutorial](https://www.google.com/search?q=python+comprehension+tutorial) en de [video](https://www.youtube.com/watch?v=3dt4OGnU5sM) is prima, maar ook de [datacamp tutorial](https://www.datacamp.com/community/tutorials/python-list-comprehension) of de [real python](https://realpython.com/list-comprehension-python/) tutorial.

### Wat opgaven

Je gaat allemaal comprehensies doen op een gedicht van annie MG Schmidt. Het staat in de cel hieronder als string in de variabele `AMGS` en daarna als lijst van woorden in `t`.

In je tests gebruik je allemaal string functies, en ook `in`. 

#### list comprehensies

1. de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)
2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)
3. de lijst van alle woorden in `AMGS` van lengte meer dan 3
4. de lijst van alle woorden in `AMGS` waar een `a` in zit
5. nu waar juist geen `a` in zit
6. nu waar als er een `a` in zit ook een `b` in zit (wat betekent dit precies he?)
7. alle woorden zonder klinker.
8. Voor elk woord zijn lengte.
9. De lijst met het hele gedicht maar nu waar elk woord met een hoofdletter begint.
10. De lijst die net zo lang is als het gedicht (dus `t`), maar dan op elke positie of het woord op die positie gelijk is aan `ik`. Dus de waardes zijn alleen maar `True` en `False`.
    * Gebruik `sum` om te berekenen hoe vaak het woord `ik` voorkomt.
11. Verzien er zelf 10 bij, en maak die ook, en geef ze aan je groepsgenoeten.

#### set comprehensies

1. De verzameling van alle unieke woorden in `AMGS`.
2. De verzameling van alle unieke woorden in `AMGS` met `ch` er in.
3. De verzameling van alle unieke woorden in `AMGS`, maar dan geschreven in hoofdletters.
4. Alle woordlengtes die in `AMGS` voorkomen.


### Dict comprehensies

1. De dict met voor elk uniek "woord" in `AMGS` zijn lengte. Dus het woord is de key en de lengte de waarde. Let op , woorden zijn dus de elementen in `t`. (antwoord beneden)
2. De dict met voor elk uniek "woord" in `AMGS` de verzameling karakters in dat woord.
3. De dict met voor elk uniek "woord" in `AMGS`  de "schoongemaakte" versie van dat woord (zonder leestekens, etc, gebruik een lange lijst als `a.replace(':','').replace(',','')`. Zet het woord ook in klein eletters.

    
### Nesting

3. De dict met voor elk uniek woord, de verzameling klinkers in dat  woord. 
    *   (antwoord beneden, probeer eerst minstens een half uur zelf)
4.  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.
    (antwoord ook beneden, niet meteen gan kijken.)
5. DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!   Hint loop over de posities. 
6. Na de vorige is deze heel makkelijk: DE dict met voor elk woord  hoe vaak het voorkomt. 


```python
AMGS='''
Ik wil niet meer, ik wil niet meer!
Ik wil geen handjes geven!
Ik wil niet zeggen elke keer:
Jawel mevrouw, jawel meneer...
nee, nooit meer in m'n leven!
Ik hou m'n handen op m'n rug
en ik zeg lekker niks terug!

Ik wil geen vieze havermout,
ik wil geen tandjes poetsen!
Ik wil lekker knoeien met het zout,
ik wil niet aardig zijn, maar stout
en van de leuning roetsen
en schipbreuk spelen in de teil
en ik wil spugen op het zeil!

En heel hard stampen in een plas
en dan m'n tong uitsteken
en morsen op m'n nieuwe jas
en ik wil overmorgen pas
weer met twee woorden spreken!
En ik wil alles wat niet mag,
de hele dag, de hele dag!

En ik wil op de kanapee
met hele vuile schoenen
en ik wil aldoor gillen: nee!
En ik wil met de melkboer mee
en dan het paardje zoenen.
En dat is alles wat ik wil
en als ze kwaad zijn, zeg ik: Bil!'''

t=AMGS.split()
len(t),t[:10]
```

```python
#Opgaven list comprehensions
#1. de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)
listcomp_1 = [w for w in t if w.startswith('i')]
#2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)
listcomp_2 = [w for w in t if w.startswith('i') or w.startswith('I')]
#3. de lijst van alle woorden in `AMGS` van lengte meer dan 3
listcomp_3 = [word for word in t if len(word) > 3]
#4. de lijst van alle woorden in `AMGS` waar een `a` in zit
listcomp_4 = [word for word in t if 'a' in word.lower()]
#5. nu waar juist geen `a` in zit
listcomp_5 = [word for word in t if 'a' not in word.lower()]
#6. nu waar als er een `a` in zit ook een `b` in zit (wat betekent dit precies he?)
listcomp_6 = [word for word in t if 'a' not in word or ('a' in word and 'b' in word)]
#7. alle woorden zonder klinker.
listcomp_7 = [word for word in t if all([letter.lower() not in 'aeiouy' for letter in word])]
#8. Voor elk woord zijn lengte.
listcomp_8 = [len(word) for word in t]
#9. De lijst met het hele gedicht maar nu waar elk woord met een hoofdletter begint.
listcomp_9 = [word[0].upper() + word[1:] for word in t]
#10. De lijst die net zo lang is als het gedicht (dus `t`), maar dan op elke positie of het woord op die positie gelijk is aan `ik`. Dus de waardes zijn alleen maar `True` en `False`.
listcomp_10 = [word.lower() == 'ik' for word in t]
print(sum(listcomp_10))

#Opgaven set comprehensions
#1. De verzameling van alle unieke woorden in `AMGS`.
setcomp_1 = {word for word in t}
#2. De verzameling van alle unieke woorden in `AMGS` met `ch` er in.
setcomp_2 = {word for word in t if 'ch' in word}
#3. De verzameling van alle unieke woorden in `AMGS`, maar dan geschreven in hoofdletters.
setcomp_3 = {word.upper() for word in t}
#4. Alle woordlengtes die in `AMGS` voorkomen.
setcomp_4 = {len(word) for word in t}

#Opgaven Dict comprehensions
#1. De dict met voor elk uniek "woord" in `AMGS` zijn lengte. Dus het woord is de key en de lengte de waarde. Let op , woorden zijn dus de elementen in `t`. (antwoord beneden)
dictcomp_1 = {word:len(word) for word in t}
#2. De dict met voor elk uniek "woord" in `AMGS` de verzameling karakters in dat woord.
dictcomp_2 = {word:{letter for letter in word} for word in t}
#3. De dict met voor elk uniek "woord" in `AMGS`  de "schoongemaakte" versie van dat woord (zonder leestekens, etc, gebruik een lange lijst als `a.replace(':','').replace(',','')`. Zet het woord ook in klein eletters.
dictcomp_3 = {word: word.replace(':','').replace(',','').replace('!','').replace('.','') for word in t}
dictcomp_3_2 = {word: ''.join([letter for letter in word if letter not in ':,!?.']) for word in t} #andere methode

#Opgaven Nested comprehensions
#3. De dict met voor elk uniek woord, de verzameling klinkers in dat  woord. 
nestcomp_3 = {word: {letter for letter in word if letter.lower() in 'aeiou'} for word in t}
#4.  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.
nestcomp_4 = {indx: {t[indx-1], t[indx+1]} for indx in range(1, len(t)-1)}
#5. DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!   Hint loop over de posities. 
nestcomp_5 = {word: {index for index, item in enumerate(t) if item == word} for word in t}
#6. Na de vorige is deze heel makkelijk: DE dict met voor elk woord  hoe vaak het voorkomt. 
nestcomp_6 = {word: t.count(word) for word in t}

```

```python
# probeer eens alleen t[0]. en dan shift return, en probeer eens wat van de suggesties
# om erachter te komen wat iets doet typ je t[0].isupper? en dan shift enter.
# Ideaal he! Dat kan bij elk object. t[0] is natuurlijk een string, dus je krijgt string functies te zien. 
t[0].isupper()
```

```python
# "in" is ook heel handig in Python

print(t[1])
'i' in t[1], 'W' in t[1], 'W' in t[1].upper()
```

```python
#de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)
[w for w in t if w.startswith('i')]
```

```python
#2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)
# Verzin minstens 2 andere manieren om exact hetzelfde te krijgen
print([w for w in t if w.startswith('i') or w.startswith('I')])
print([w for w in t if w[0] == ('i') or w[0] == ('I')])
print([w for w in t if w.lower()[0] == ('i')])
```

```python
#  De dict met voor elk uniek woord in `AMGS` zijn lengte.

{w:len(w) for w in t}
```

```python
# De dict met voor elk uniek woord, de verzameling klinkers in dat  woord.

{w:{l.lower() for l in w if l in 'aeiouAEIOU'} 
    for w in set(t)} # verbeter hem zelf zodat hoofdletters ook mee worden genomen

```

```python
#  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.

{i:{t[i-1],t[i+1]} for i in range(1,len(t)-1)}
```

<!-- grade_id: TRICKY, points: 0, grade: false, solution: true -->
```python
#  DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!    

l= {word: {index for index, item in enumerate(t) if item == word} for word in t}
print(l)
```

```python

```

```python

```

```python

```