# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 5

```python
import random
import pandas as pd
import networkx as nx
import operator
import os
import seaborn as sns
import numpy as np
import matplotlib.pyplot as plt
import scipy.stats as stats
from collections import defaultdict, Counter
from nose.tools import assert_equal
from numpy.testing import assert_almost_equal
%matplotlib inline
```

# First course Chap 3, exercises 3.1-3.10

Maak ze allemaal netjes in de cel hieronder.

<!-- grade_id: VraagChap3, points: 10, grade: true, solution: true -->
**3.1** Done

**3.2** Average Degree = 200/100 = 2

**3.3**
1. Is this graph connected? **No**
2. What is the mode (most frequent value) of the node degree distribution? **One**
3. How many nodes are in the largest clique? **Four**
4. Would you expect this graph to have any hubs? **No**

**3.4** max(G.nodes, key=nx.degree_centrality(G).get) and G.degree(max(nx.degree_centrality(G)))

**3.5** B

**3.6** D

**3.7** B

**3.8** B

**3.9** For closeness centrality to be well-defined, the network must be strongly connected, ensuring that there is a path between any pair of nodes.

**3.10**
1. The node with the highest degree is not the one with largest closeness

    *The node with the largest closeness is the **<font color='red'>red</font>** node*
    
    *The ones with the highest degree is the **<font color='black'>black</font>** one*
    
2. The node with the highest betweenness is not the one with largest closeness

    *The nodes with the largest betweenness are the **<font color='black'>black</font>** ones*
    
    *The ones with the largest closeness are the **<font color='red'>red</font>** ones*

```python
# Question 3.10
G1 = nx.Graph()

G1.add_nodes_from([1,2,3,4,5,6,7])
G1.add_edges_from([(1,2),(2,3),(3,4),(4,5),(5,6),(5,7),])

color_mapG1 = ['blue','blue','blue','red','black','blue','blue',]

nx.draw(G1, node_color=color_mapG1, with_labels=True)
```

```python
# Question 3.10
G2 = nx.Graph()

G2.add_nodes_from([1,2,3,4,5,6,7,8])
G2.add_edges_from([(1,2),(1,3),(2,3),(3,4),(3,5),(4,6),(5,6),(6,7),(6,8),(7,8)])

color_mapG2 = ['blue','blue','black','red','red','black','blue','blue']

nx.draw(G2, node_color=color_mapG2, with_labels=True)
```

# First course Chap 2, exercise 2.41

```python
# load the data; vervang de fname naar jouw pad

fname="/Users/rvanheusden/Desktop/courses2023/NetworkScience/Teacher/FirstCourseInNetworkScience/FirstCourseNetworkScience/datasets/socfb-Northwestern25/socfb-Northwestern25.edges.gz"
G= nx.Graph()
G = nx.read_edgelist(fname)
print(len(G))  # Do not try to draw this network
```

```python
# probeer maar eens met 1000. Dat wil je niet.
sample= np.random.choice(np.array(G), 100, replace=False)
nx.draw_kamada_kawai(G.subgraph(sample))
```

<!-- grade_id: 241a, points: 0, grade: false, solution: true -->
```python
temp_graph = temp_graph = G.subgraph(sample)

hoeveel_knopen= nx.number_of_nodes(temp_graph)
hoeveel_lijnen=nx.number_of_edges(temp_graph)
dichtheid=nx.density(temp_graph)
connected= nx.is_connected(temp_graph)
how_many_pairs_of_nodes= len(temp_graph.edges())


hoeveel_knopen, hoeveel_lijnen, dichtheid, connected, how_many_pairs_of_nodes

```

<!-- grade_id: 241t1, points: 1, grade: true, solution: false -->
```python

```

# Vraag 2.41 deelvraag 4 

* Zorg dat je eerst test of `G` verbonden is, en zo niet bereken het dan over de giant component

<!-- grade_id: 241a2, points: 0, grade: false, solution: true -->
```python
#random.sample(G.nodes,2)

def shortest_path_length_sample(G,samplesize=1000):
    '''Return a list with samplesize shortest path lengths'''
    new_list = []
    H = G
    if nx.is_connected(G) == False:
        nodes_in_giant_comp = list(nx.connected_components(H))[0]
        H = nx.subgraph(H, nodes_in_giant_comp)
    for n in range(0, samplesize):
        random_nodes = random.sample(H.nodes,2)
        new_list.append(nx.shortest_path_length(G, random_nodes[0], random_nodes[1]))
    return new_list

# run 
sample_lengths= shortest_path_length_sample(G,100)   


# gemiddelde  

print("Gemiddelde=",round(np.mean(sample_lengths),1))
      
S= pd.Series(sample_lengths)
print(S.describe())
sns.displot(S);
```

<!-- grade_id: 241t2, points: 1, grade: true, solution: false -->
```python
L= shortest_path_length_sample(G,10)
assert isinstance(L,list)
assert isinstance(np.mean(L), float)

```

# En schat nu de diameter

<!-- grade_id: 241a3, points: 0, grade: false, solution: true -->
```python
diameter= 3

 

print("Diameter",diameter)
# run 
sample_lengths= shortest_path_length_sample(G,100)   
      
S= pd.Series(sample_lengths)
print(S.describe())
sns.displot(S);
```

# En nu de clustering coefficient

* Probeer eerst eens met networkx, maar dan op de karate club graaf.
* Het gemiddelde kan `pandas` op `numpy` mooi voor je uitrekenen.
    * We late zien hoe dat werkt op een willekeurige dict.
    * Maak het zelf af. Kijk ook eens naar dat histogram. 
* Zijn degrees normaal verdeeld?

```python
# Voorbeeld: hoe vaak komt elke degree voor in de karateclub

H= nx.karate_club_graph()

alle_degrees=pd.Series(dict(H.degree()))
print(alle_degrees.describe())  # sleutel waardes van de degree verdeling
print("gemiddelde is ",alle_degrees.describe()['mean'])
sns.displot(alle_degrees, kde=True);  # hoe vaak komt elke degree voor
```

<!-- grade_id: cca1, points: 0, grade: false, solution: true -->
```python
%%time

H = nx.karate_club_graph()

alle_clustering = pd.Series(nx.clustering(H))
print(alle_clustering.describe())  # sleutel waardes van de degree verdeling

gemiddelde = alle_clustering.describe()['mean']

print(gemiddelde)
```

# Assortative or not?

1. Maak eerst de functie `my_avg_nn_degree(Network,degree)` af, die dus hetzelfde doet als `nx.k_nearest_neighbors(H)[k]`.  Nu snap je wat ze daarmee bedoelen. In het boek op bladzijde 38-39 staat het best cryptisch, vond ik.
2. Maak dan  nu de opgave uit het boek.

<!-- grade_id: assa1, points: 0, grade: false, solution: true -->
```python
 
#overgeslagen

def my_avg_nn_degree(Network,degree):
    '''Geeft de gemiddelde degree van de buren, gemiddeld over alle knopen van degree degree.
    In het Engels: 
    It gives the average of the average_neighbor_degree of all nodes with the degree degree.
    average_neighbor_degree(k) for k a node is the average of the degrees of the neighbors of k.'''
    WRITE YOUR CODE HERE
    
#test, ER ZIJN IN NETWORKX DUS HEEL VEEL MANIEREN OM DIT UIT TE REKENEN, 
# hier kijken we of die allemaal gelijk zijn, voor elke voorkomende degree in de karateclub graaf

H= nx.karate_club_graph()
for k in {k[1] for k in H.degree()}:
    print(k, 
          
          nx.k_nearest_neighbors(H)[k] , 
          my_avg_nn_degree(H,k),
          nx.algorithms.assortativity.average_degree_connectivity(H)[k]
         )

    
# mooiere test; kijk eens goed, wat gaat er nu mis? 

all(nx.k_nearest_neighbors(H)[k[1]]== my_avg_nn_degree(H,k[1]) for k in H.degree())
```

<!-- grade_id: assa2, points: 0, grade: false, solution: true -->
```python
import scipy.stats as stats

assortativity1=0 # volgens de blauwe box op blz 38-39
assortativity2=0 # en de tweede manier
WRITE YOUR CODE HERE
assortativity1,assortativity2
```

<!-- grade_id: ass1t, points: 1, grade: true, solution: false -->
```python
assert isinstance(my_avg_nn_degree(nx.karate_club_graph(),1), float)

```

<!-- grade_id: ass2t, points: 1, grade: true, solution: false -->
```python
assert isinstance(assortativity1, float)
assert isinstance(assortativity2, float)
```

# Deel 1: Easley & Kleinberg Chapter 5

## Vraag 1
Maak opdracht [5.6.1 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch05.pdf). Geef je antwoord in de cell hieronder


<!-- grade_id: Vraag1_antwoord, points: 1, grade: true, solution: true -->
Yes, this network is balanced, since everybody knows eachother so the triangle consists of three positive edges and two people of the same village have the same enemy so the triangle consists of two negative edges and one positive edge. And both forms of the triangles are balanced so the whole network is balanced too.

## Vraag 2
Maak opdracht [5.6.2 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch05.pdf)

Geef als tuple voor `pos_edge` en `neg_edge`

Dus `pos_edge = (balanced, unbalanced)`, waar balanced en unbalanced het aantal triangles zijn

<!-- grade_id: Vraag2, points: 0, grade: false, solution: true -->
```python
pos_edge = (1,2)
neg_edge = (2,1)


```

<!-- grade_id: Vraag2_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(pos_edge), tuple)
assert_equal(type(neg_edge), tuple)
assert_equal(type(pos_edge[0]), int)
assert_equal(sum(pos_edge), 3, "Lees de vraag goed!")
```

## Vraag 3

Maak opdracht [5.6.3 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch05.pdf). Geef je antwoord in de cell hieronder

<!-- grade_id: Vraag3_antwoord, points: 1, grade: true, solution: true -->
a) Node D could become friends with node B and C and enemies with node A. In this way al the possible triangles with node D will have a -++ structure. This creates a balanced network. 

b) Node D could become friends with node A, B and C. In this way al the possible triangles with node D have an +++ structure. This creates a balanced network. Another option is for node D to become enemies with nodes A,B and C. In this way al the possible triangles with node D have an --+ structure. This also creates a balanced network. So for this situation there are 2 possible options.

c) Node X should become friends with all the nodes if the network contains a ++- structure. By this way node X will not get in an unbalanced triangle. If the network structure contains only negative relationships (---) node X should become enemies with one node and enemies with the other nodes. This makes a +-- structure which causes the triangles with node X to be balanced.

## Vraag 4

Maak opdracht [5.6.4 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch05.pdf). Geef je antwoord in de cell hieronder

<!-- grade_id: Vraag4_antwoord, points: 1, grade: true, solution: true -->
This network would be structurally unbalanced. A farmer is friends with a farmer that lives at most 20 miles away. Each farmer
lives on a tract of land that occupies a 1-mile stretch of the river bank. Farmer A is friends with all the farmers in a 20 mile range and is an enemy with all the farmers outside the 20 mile range. Farmer A is friends with farmer B because he lives in the 20 mile range but farmer A would be enemies with farmer C because farmer C lies outside this range. Farmer B on the other hand would be friends with farmer C. This makes a triangle with the ++- structure. This is an unbalanced network.

# Deel 2: Moties in de Tweede Kamer

```python
def open_file():
    """Loads the graphml"""
    if 'motieslist.csv' in os.listdir():
        return pd.read_csv(
            'motieslist.csv', encoding='utf-8', compression='gzip')
    elif os.path.exists('../../data/Week5'):
        return pd.read_csv(
            '../../data/Week5/motieslist.csv', encoding='utf-8', compression='gzip')
    elif os.path.exists('../../../data/Week5'):
        return pd.read_csv(
            '../../../data/Week5/motieslist.csv',
            encoding='utf-8',
            compression='gzip')


MOTnetwerk = open_file()
MOTnetwerk.datum=MOTnetwerk.datum.astype(int)
MOTnetwerk.head()
```

```python
motie_dict= MOTnetwerk.T.to_dict()
print("We hebben %i moties! " %len(set(motie_dict[key]['URL naar motie'] for key in motie_dict)))

# vergelijk 3 keer hetzelfde
print(MOTnetwerk.loc[0])
print(motie_dict[1])
for attribuut in motie_dict[1].keys():
    print(attribuut,motie_dict[1][attribuut])
```

## Vraag 2

We hebben nu alle moties uit het csv bestand gehaald, en er een dict van dicts van gemaakt.
Voor deze opdracht zijn we benieuwd naar de Tweede Kamer tijdens het missionaire [kabinet Rutte III](https://nl.wikipedia.org/wiki/Kabinet-Rutte_III).   Beperk dus het bestand tot alle moties vanaf (en inclusief) de dag dat dat kabinet beëdigd is.

Noem dit `MotiesRutteIII`, met dezelfde structuur als `motie_dict` hierboven. Eigenlijk is `MotiesRutteIII` een *deelverzameling* van `motie_dict`.

**Check** er zitten 4498 moties in RutteIII. Zoek de juiste datum dus op.

<!-- grade_id: Vraag2M, points: 0, grade: false, solution: true -->
```python
MotiesRutteIII = {}
for key, value in motie_dict.items():
    datum = value['datum']
    if datum >= 20171026:
        MotiesRutteIII[key] = value

print("We hebben %s moties tijdens Rutte III!" % 
      len(set(motie_dict[key]['URL naar motie'] for key in MotiesRutteIII)))
print("een voorbeeld:")
MotiesRutteIII[2269]
```

<!-- grade_id: Vraag2M_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(MotiesRutteIII), dict)
 
```

## Vraag 3

Maak de verzameling van alle indieners (dus niet medeindieners!) waarin elke indiener maar 1 keer voorkomt gebaseerd op `MotiesRutteIII`. Noem deze `indieners`

<!-- grade_id: Vraag3M, points: 0, grade: false, solution: true -->
```python
indieners = set(MotiesRutteIII[x]['indiener'] for x in MotiesRutteIII)

len(indieners), list(indieners)[:5]
```

<!-- grade_id: Vraag3M_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(indieners), set)
for indiener in indieners:
    assert_equal(type(indiener), str)
```

## Vraag 4
Maak een dict `IndienersDict` waarbij de key indiener X is van moties, en de value
de partij is van X. Doe hetzelfde voor medeindieners in de dict `MedeindienersDict` en een dict `CombinedDienersDict` met zowel indieners als medeindieners. Dit is weer gebaseerd op de moties gedurende Rutte III.

Bekijk hoeveel je er van elk hebt en hoeveel in totaal. Gebruik een dict-comprehension!

Deel van het antwoord van CombinedDienersDict:

`{'A. Pechtold': 'D66',
 'A.C.L. (Arno) Rutte': 'VVD',
 'Achraf Bouali': 'D66',
 'Agnes Mulder': 'CDA',
 'Albert van den Bosch': 'VVD',
 'Alexander Kops': 'PVV',
 'André Bosman': 'VVD',
 'Anne Kuik': 'CDA',
 'Anne Mulder': 'VVD',...}`
 
 **Let op** Dicts zijn ongeordend. Jouw antwoord hoeft er dus niet zo uit te zien als hierboven als je het print.

<!-- grade_id: Vraag4M, points: 0, grade: false, solution: true -->
```python
IndienersDict = {MotiesRutteIII[x]['indiener']: MotiesRutteIII[x]['Partij Indiener'] for x in MotiesRutteIII}
MedeindienersDict = {MotiesRutteIII[x]['Medeindiener']: MotiesRutteIII[x]['Partij Medeindiener'] for x in MotiesRutteIII}
CombinedDienersDict = {**IndienersDict, **MedeindienersDict}

len(IndienersDict), len(MedeindienersDict), len(CombinedDienersDict)
```

<!-- grade_id: Vraag4M_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(IndienersDict), dict)
assert_equal(type(MedeindienersDict), dict)
assert_equal(type(CombinedDienersDict), dict)
```

## Vraag 5
Maak een networkx netwerk `G`, waarin zowel alle indieners als medeindieners een knoop zijn, en waarbij je aan elke  knoop toevoegt  van welke partij diegene is. Je voegt dus een node -attribuut toe. Noem dit attribuut `partij`.

We willen een <em>gericht</em> netwerk maken, dus gebruik `G = nx.DiGraph()`. 

<!-- grade_id: Vraag5, points: 0, grade: false, solution: true -->
```python
G = nx.DiGraph()
G.add_nodes_from([(node, {'partij':partij}) for (node, partij) in CombinedDienersDict.items()])

#([(node, {'partij':partij}) for (node, partij) in CombinedDienersDict.items()])
print('De partij van Suzanne Kröger is', G.nodes['Suzanne Kröger']['partij'])
nx.draw(G, with_labels=True)
plt.show()
```

<!-- grade_id: Vraag5_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(nx.is_directed(G), True)
assert_equal(type(nx.get_node_attributes(G, 'partij')), dict)
```

## Vraag 6
Maak een lijst `InEnMedeIndiener` met paren bestaande uit indieners en medeindieners van dezelfde motie.
Het eerste element van elk paartje is dus de indiener. WE gebruiken een lijst (en geen verzameling) want we zijn ook geinterreseerd in hoe vaak mensen samen indienen.

Je lijst heeft dus de vorm:

`[('M.L. Thieme', 'Jesse Klaver'),
 ('Lodewijk Asscher', 'E.G.M. Roemer'),
 ('Lodewijk Asscher', 'Jesse Klaver'),
 ('Lodewijk Asscher', 'Henk Krol'),
 ('Lodewijk Asscher', 'M.L. Thieme')....]`

<!-- grade_id: Vraag6, points: 0, grade: false, solution: true -->
```python
InEnMedeIndiener = [(MotiesRutteIII[x]['indiener'], MotiesRutteIII[x]['Medeindiener']) for x in MotiesRutteIII]

InEnMedeIndiener[:5]
```

<!-- grade_id: Vraag6_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(InEnMedeIndiener), list)
```

## Vraag 7
Maak een lijst `InEnMedeIndienerCounts` met tuples van de vorm:

`[('M.L. Thieme', 'Jesse Klaver', 9),
 ('Lodewijk Asscher', 'E.G.M. Roemer', 3),
 ('Lodewijk Asscher', 'Jesse Klaver', 20),
 ('Lodewijk Asscher', 'Henk Krol', 6),
 ('Lodewijk Asscher', 'M.L. Thieme', 6)...]`

waarin elk tuple het volgende betekent:
    (indiener, medeindiener, "hoe vaak indiener een motie heeft ingediend met medeindiener")

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
from collections import Counter  # Handig!

InEnMedeIndienerCounts = [(indiener, medeindiener, count) for (indiener, medeindiener), count in Counter(InEnMedeIndiener).items()]

InEnMedeIndienerCounts[:5]
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(InEnMedeIndienerCounts), list)
```

## Vraag 8
Voeg aan je netwerk   alle  weighted edges (A,B) toe waarbij A een indiener is en
    B een medeindiener en  waarin
    het gewicht aangeeft hoe vaak indiener A een motie heeft ingediend met medeindiener B.
    Gebruik edge attributes hiervoor en het speciale attribute `weight`.
       **Je gebruikt natuurlijk je mooie lijst uit de vorige vraag.**

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
G.add_weighted_edges_from([(indiener, medeindiener, count) for indiener, medeindiener, count in InEnMedeIndienerCounts])

# 3 testjes achter elkaar

list(G.edges(data=True))[:5], G['Suzanne Kröger']['Cem Laçin'], G['Cem Laçin']['Suzanne Kröger']
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(nx.is_directed(G), True)
assert_equal(type(nx.get_edge_attributes(G, 'weight')), dict)
```

```python
# als dit niet werkt, gaat er waarschijnlijk wat fout

def draw_Rutte(G):
    partijdict = defaultdict(list)
    for key, values in nx.get_node_attributes(G, 'partij').items():
        partijdict[values].append(key)

    plt.figure(figsize=(25, 25))

    node_list = []
    for partij in partijdict.keys():
        node_list += list(partijdict[partij])

    pos = {}
    for i, node in enumerate(node_list):
        theta = 2.0 * np.pi * i / len(node_list)
        pos[node] = (np.cos(theta), np.sin(theta))

    color = sns.color_palette("hls", len(partijdict.keys()))
    for i, j in enumerate(partijdict.keys()):
        c = np.array(len(partijdict[j]) * color[i]).reshape((len(partijdict[j]),3))
        nx.draw_networkx_nodes(
            G, pos, nodelist=partijdict[j], node_color=c, label=j)

    max_weight = max(nx.get_edge_attributes(G, 'weight').values())
    weights = [G[u][v]['weight'] / max_weight for u, v in G.edges()]
    nx.draw_networkx_edges(G, pos, width=weights)

    plt.legend()
    plt.axis('off')
    plt.show()


draw_Rutte(G)
```

## Vraag 9 

* Maak een verzameling `PVV` met alle knopen uit `G` die lid zijn van de PVV.
* Gebruik nu `draw_Rutte` om het samenwerk verband tussen PVVers te tekenen.
* Kijk ook eens naar andere partijen. Valt je iets op?

**HINT** Maak eens de verzameling van alle partijen die voorkomen in netwerk `G`. Dat is 1 set comprehensie (waarbij je niet meer naar `motie_dict` gaat).

We willen nu voor elke partij weten wat de _dichtheid_ (in het Engels _density_) is van de deelgraaf van `G` geinduceerd door de leden van die partij, deze dictionary noemen we `dichtheid_per_partij`.
* Dat maak je gewoon heel makkelijk met 1 dict comprehensie in 1 regel, met daarin nog een geneste set comprehensie.
* Maar goed het mag ook anders hoor.
* Je antwoord ziet er zo uit (of in een andere volgorde natuurlijk):

```
dichtheid_per_partij = {'SGP': 0.16666666666666666, 'DENK': 0.8333333333333334, 'PvdA': 0.13333333333333333,
 ```

<!-- grade_id: vraag9, points: 0, grade: false, solution: true -->
```python
PVV = {x for x in G if G.nodes[x]['partij'] == 'PVV'}
dichtheid_per_partij = dict({x: nx.density(G.subgraph(list({y for y in G if G.nodes[y]['partij'] == x}))) for x in set(G.nodes[y]['partij'] for y in G)})

draw_Rutte(G.subgraph(PVV))
# het antwoord even mooi printen
sorted(dichtheid_per_partij.items(), key=lambda kv: kv[1])
```

<!-- grade_id: vraag9_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(PVV), set)
assert_equal(type(dichtheid_per_partij), dict)

```

## Vraag 10
Bereken in- en outdegree voor elk knoop in het Rutte netwerk `G`.
Maak hiervoor de twee lijsten `indegree` en `outdegree`, met daarin de tuples met als eerste waarde de degree en als tweede waarde de naam. Sorteer aflopend op de degrees, en bij gelijke degree omgekeerd alfabetisch op naam.

Voorbeeld voor indegree:

`[(52, 'Chris Stoffer'),
  (46, 'Roelof Bisschop'),
  (45, 'Eppo Bruins'),
  (45, 'Carla Dik-Faber'),
  (44, 'Kees van der Staaij')...]`

Op https://data.kro-ncrv.nl/kamerleden-2018/ kun je zien dat een aantal van hen ook bij de meest succesvolle kamerleden (hoogste percentage aangenomen moties en amendementen) van deze kabinetsperiode horen.

**Hint:** Kijk eens naar de code in de volgende vraag..... Dat helpt he?

<!-- grade_id: Vraag10, points: 0, grade: false, solution: true -->
```python
indegree = sorted([(G.in_degree(x), x) for x in G], reverse = True)
outdegree = sorted([(G.out_degree(x), x) for x in G], reverse = True)

indegree[:5], outdegree[:5]
```

<!-- grade_id: Vraag10_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(indegree), list)
assert_equal(type(indegree[0]), tuple)
assert_equal(type(outdegree), list)
assert_equal(type(indegree[0]), tuple)

```

## Vraag 11

Beschouw het Rutte netwerk `G` nu als <strong>een ongerichte graaf</strong> en bepaal de betweenness centrality van elke knoop.

Maak een dictionary `bcDict` met als key de knoop en als value de betweenness centrality.

Voorbeeld: `{'A. Pechtold': 0.0,
 'A.C.L. (Arno) Rutte': 0.0008961172303149877,
 'Achraf Bouali': 0.0009292316373274312,
 'Agnes Mulder': 0.005956492682257983,
 'Albert van den Bosch': 0.0005094599555776784,
 'Alexander Kops': 0.006054635024810758,
 'André Bosman': 0.003933997703294198,
 'Anne Kuik': 0.007195012823896772,
 'Anne Mulder': 6.656473869236166e-05,..}`

<!-- grade_id: Vraag11, points: 0, grade: false, solution: true -->
```python
bcDict = {x:nx.betweenness_centrality(G)[x] for x in G}



#even kijken naar wie de hoogste scores hebben
sorted(bcDict.items(), key=operator.itemgetter(1), reverse=True)[:5]

```

<!-- grade_id: Vraag11_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(bcDict), dict)
assert_equal(type(list(bcDict.keys())[0]), str)
assert_equal(type(list(bcDict.values())[0]), float)

```

# Theorie vragen

## Vraag 17

Given a complete   undirected graph G with n vertices. 

1. Argue
that a given vertex u can lie on XX  shortest paths connecting
                two other (distinct) vertices.  Find XX yourself.  

2. On how many paths without duplicates does u lie? Give your argument. Hint: this is not easy! You must count very precisely.
    * Give the argument, and then program it as well. (You get half a point for both)
            

<!-- grade_id: v17a, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: v17aa, points: 0, grade: false, solution: true -->
```python
WRITE YOUR CODE HERE

def vraag17(n):
    WRITE YOUR CODE HERE
vraag17(20), vraag17(3), vraag17(4), vraag17(30), type(vraag17(3))
```

<!-- grade_id: v17aT, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(vraag17(3), np.int)


```

## Vraag 18 (2pt)

1. Geef in het Nederlands de definitie van een  [ongericht bipartiet netwerk](https://en.wikipedia.org/wiki/Bipartite_graph).
    * Je definitie moet zo precies zijn dat als je er een netwerk G "aan geeft" dat de defnitie dan kan bepalen of het bipartiet is of niet. 
    * Je definitie is dus al bijna een _algorithme_ of een _programma_.

2. Gebruik je definitie   om een echt programma te maken. Het probeert gewoon elke verdeling.
    * Van networkx mag je alleen `G.nodes, G.edges, G.subgraph` gebruiken.
    * Hint: <https://www.geeksforgeeks.org/python-program-to-get-all-subsets-of-given-size-of-a-set/>
    * Hoeveel checks moet je uitvoeren bij een netwerk van $n$ knopen?
    * Hoeveel bij  netwerken van 3,5,10,20,30 knopen? (je kan dit zonder rekenmachine).

<!-- grade_id: v18a1, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: v18a2, points: 0, grade: false, solution: true -->
```python
%%time
WRITE YOUR CODE HERE
def bipartiet(G):
    '''Returns True iff G is bipartite'''
    WRITE YOUR CODE HERE

G = nx.house_graph()
print(bipartiet(G))
#nx.draw_kamada_kawai(G, with_labels=True)
H= nx.Graph()
H.add_edges_from([(i,i+1) for i in range(10)])
nx.draw_kamada_kawai(H, with_labels=True)

print(bipartiet(H))
```

<!-- grade_id: v18a2t, points: 1, grade: true, solution: false -->
```python

```