# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


antwoorden ch2: <br>
2.2: True<br>
2.3: True, elke node heeft een parent en een child<br>
2.4: connected network: N-1 , niet connected network: 0, nodes hoeven niet verbonden te zijn met elkaar. <br>
2.5: True<br>
2.6: True<br>
2.7: False<br>
2.8: cycle: [(D, E), (E, F), (F, D)], weakly connected<br>
2.9: False<br>
2.10: 2<br>
2.11: ja, link gaan nu in plaats van een kant twee kanten op, dus blijft het netwerk connected.<br>
2.12: C<br>
2.13: B<br>
2.14:<br>
in-degree D: 11<br>
out-degree C: 3<br>
2.15: 1 <br>
2.16: B <br>
2.17: [1, 2, 4, 6] (0.8 + 0.9 + 0.3 = 2 en 0.6 + 0.5 + 0.95 = 2.05)<br>
2.18: [1, 3, 5, 6] (1/0.8 + 1/0.9 + 1/0.3 = 5.694... en 1/0.6 + 1/0.5 + 1/0.95 = 4.719...)<br> 
2.19: C (de korste afstand tussen de twee nodes die het meest uit elkaar liggen)<br>
2.20: A (er zijn bijna geen buren van knopen verbonden met elkaar)<br>
2.21: Nee, een hogere diameter en clustering coefficient<br>
2.22: A<br>
2.23: 5<br>
2.24: 3 <br>
2.25: False <br>
2.26: C <br>
2.27: True<br>
2.28: D <br>
2.32: 1 <br>
2.33: 0 <br>
2.34: 0.3 (2*3)/(4*5)<br>
2.35: {(1, 2): 1, (1, 3): 2, (1, 4): 2, (1, 8): 3, (1, 5): 3, (1, 6): 4, (1, 7): 4, (2, 3): 1, (2, 4): 1, (2, 8): 2,
       (2, 5): 2, (2, 6): 3, (2, 7): 3, (3, 4): 1, (3, 8): 2, (3, 5): 2, (3, 6): 3, (3, 7): 3, (4, 8): 1, (4, 5): 1,
       (4, 6): 2, (4, 7): 2, (5, 8): 2, (5, 6): 1, (5, 7): 1, (6, 8): 3, (6, 7): 2, (7, 8): 3, (9,10): 1}<br>
2.36: {11: 0, 1: 0, 2: 0.3333333333333333, 3: 1.0, 4: 0.16666666666666666, 8: 0, 5: 0, 6: 0, 7: 0, 10: 0, 9: 0}<br>
2.37: {('d', 'e'): 1, ('d', 'c'): 2, ('d', 'b'): 2, ('d', 'g'): 3, ('d', 'f'): 4, ('d', 'a'): 4, ('d', 'h'): 3,
       ('e', 'c'): 1, ('e', 'b'): 1, ('e', 'g'): 2, ('e', 'f'): 3, ('e', 'a'): 3, ('e', 'h'): 2, ('c', 'b'): 1,
       ('c', 'g'): 1, ('c', 'f'): 2, ('c', 'a'): 2, ('c', 'h'): 2, ('b', 'g'): 1, ('b', 'f'): 2, ('b', 'a'): 2,
       ('b', 'h'): 1, ('g', 'f'): 1, ('g', 'a'): 1, ('g', 'h'): 2, ('f', 'a'): 1, ('f', 'h'): 3, ('a', 'h'): 3}<br>
2.38: {'d': 0, 'e': 0.3333333333333333, 'c': 0.6666666666666666, 'b': 0.3333333333333333, 'g': 0.3333333333333333,
       'f': 1.0, 'a': 1.0, 'h': 0}<br>
2.40: A<br>






```python
# 2.8
D = nx.DiGraph()
D.add_edge('A', 'B', weight=1)
D.add_edge('B', 'C', weight=2)
D.add_edge('D', 'B', weight=1)
D.add_edge('D', 'E', weight=1)
D.add_edge('E', 'F', weight=1)
D.add_edge('F', 'A', weight=2)
D.add_edge('F', 'B', weight=1)
D.add_edge('F', 'C', weight=3)
D.add_edge('F', 'D', weight=2)
D.add_edge('F', 'E', weight=1)
nx.draw(D, with_labels=True)

nx.find_cycle(D)
nx.is_strongly_connected(D)
nx.is_weakly_connected(D)
```

```python
#2.9
G = nx.Graph()
G.add_edges_from([('A', 'B'), ('B','C'), ('D', 'B'), ('D', 'E'), ('E', 'F'), ('A', 'F'), ('B', 'F'), ('F', 'C'), ('F', 'D'),
                 ('F', 'E')])
nx.draw(G, with_labels=True)
nx.is_tree(G)
```

```python
# 2.10
nx.diameter(G)
```

```python
# 2.11
nx.is_connected(G)
```

```python
# 2.13
W = nx.DiGraph()
W.add_edge('A', 'B', weight=2)
W.add_edge('B', 'F', weight=11)
W.add_edge('D', 'F', weight=2)
W.add_edge('B', 'D', weight=9)
W.add_edge('E', 'D', weight=2)
W.add_edge('C', 'E', weight=3)
W.add_edge('B', 'C', weight=1)
W.add_edge('A', 'C', weight=6)

nx.draw(W, with_labels=True)

nx.is_weakly_connected(W)
```

```python
#2.14
W.in_degree('D', weight='weight')
W.out_degree('C', weight='weight')
```

```python
# 2.15
list(nx.strongly_connected_components(W))
```

```python
# 2.16
W2 = nx.DiGraph()
W2.add_edge(1, 2, weight=0.8)
W2.add_edge(2, 4, weight=0.9)
W2.add_edge(4, 6, weight=0.3)
W2.add_edge(4, 5, weight=0.8)
W2.add_edge(5, 6, weight=0.95)
W2.add_edge(3, 5, weight=0.5)
W2.add_edge(1, 2, weight=0.8)
W2.add_edge(1, 3, weight=0.6)

nx.draw(W2, with_labels=True)

nx.is_weakly_connected(W2)

```

```python
# 2.17
nx.dijkstra_path(W2, 1, 6)
```

```python
# 2.22
D2 = nx.DiGraph()
D2.add_edges_from([(1,2), (2,5), (5,7), (7,6), (6,3), (3,1), (1,4), (4,3), (4,5), (2,4), (4,6)])

nx.draw(D2, with_labels=True)

nx.is_strongly_connected(D2)
```

```python
#2.23
nx.diameter(D2)
```

```python
# 2.24
G2 = nx.Graph()
G2 = D2.to_undirected()

nx.draw(G2)

nx.diameter(G2)
```

```python
# 2.26
def diameters(G):
    shortest_paths = {}
    shortest_paths = {(node1, node2): nx.shortest_path_length(G, node1, node2) 
                      for node1 in G.nodes() for node2 in G.nodes() if node1 != node2}
    return max(shortest_paths.values())

diameters(G2)
```

```python
#2.34
G3 = nx.Graph()
G3.add_edges_from([('ego', 1), ('ego', 2), ('ego', 3), ('ego', 4), ('ego', 5), (1,3), (3,4), (5,4)])

nx.draw(G3, with_labels=True)
nx.clustering(G3, 'ego')
```

```python
# 2.35
#node 11 en edge (9,10) weggelaten omdat anders code niet werkt
# handmatig (9,10): 1 toevoegen
G4 = nx.Graph()
G4.add_edges_from([(1,2), (2,3), (3,4), (2,4), (4,8), (4,5), (6,5), (5,7)])

nx.draw(G4, with_labels=True)

shortest_paths1 = {}
shortest_paths1 = {(node1, node2): nx.shortest_path_length(G4, node1, node2) 
                      for node1 in G4.nodes() for node2 in G4.nodes() if node1 != node2 and node1 > node2}
shortest_paths1
```

```python
#2.36
G5 = nx.Graph()
G5.add_node(11)
G5.add_edges_from([(1,2), (2,3), (3,4), (2,4), (4,8), (4,5), (6,5), (5,7), (10,9)])

clustering_coefficient = {node: nx.clustering(G5, node) for node in G5.nodes()}
clustering_coefficient
```

```python
# 2.37
G6 = nx.Graph()
G6.add_edges_from([('d', 'e'), ('e', 'c'), ('c', 'b'), ('b', 'e'), ('c', 'g'), ('b', 'g'), 
                   ('g', 'f'), ('f', 'a'), ('g', 'a'), ('b', 'h')])

nx.draw(G6, with_labels=True)

shortest_paths1 = {}
for node1 in G6.nodes():
    for node2 in G6.nodes():
        if node1 != node2:
            length = nx.shortest_path_length(G6, node1, node2)
            if (node2, node1) not in shortest_paths1.keys():
                shortest_paths1[(node1, node2)] = length
shortest_paths1
```

```python
# 2.38
clustering_coefficient = {node: nx.clustering(G6, node) for node in G6.nodes()}
clustering_coefficient
```

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
strongly_connected_components
```

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    return max(nx.strongly_connected_components(G), key=len)

def find_IN(G):
    scc = find_SCC(G)
    scc_subgraph = G.subgraph(scc)
    return {node1 for node1 in G.nodes() for node2 in scc_subgraph.nodes() 
            if nx.has_path(G, node1, node2) and node1 not in scc_subgraph.nodes()}
    
def find_OUT(G):
    scc = find_SCC(G)
    scc_subgraph = G.subgraph(scc)
    return {node1 for node1 in G.nodes() for node2 in scc_subgraph.nodes() 
            if nx.has_path(G, node2, node1) and node1 not in scc_subgraph.nodes()}
    
# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

```python
{node: nx.clustering(G, node) for node in G.nodes}

```

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
def clusteringcoefficient(G):
    clustering_coeffs = {}
    for node in G.nodes():
        neighbors = list(G.neighbors(node))
        degree = len(neighbors)
        if degree <= 1:
            clustering_coeffs[node] = 0.0
        else:
            numerator = 0
            for i in range(degree):
                for j in range(i+1, degree):
                    neighbor1 = neighbors[i]
                    neighbor2 = neighbors[j]
                    if G.has_edge(neighbor1, neighbor2):
                        numerator += 1
            clustering_coeffs[node] = 2 * numerator / (degree * (degree-1))
    return clustering_coeffs
    

# even er een mooi plotje van maken
pd.Series( clusteringcoefficient(G) ).sort_values().plot(kind='barh',figsize=(8,8))

clusteringcoefficient(G)
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
def cc_nice(G, v):
    neighbors = [node for node in G.nodes() if G.has_edge(v, node) and node != v]
    degree = len(neighbors)
    if degree <= 1:
        return 0
    else:
        numerator = 0
        for i in range(degree):
            for j in range(i+1, degree):
                if G.has_edge(neighbors[i], neighbors[j]):
                    numerator += 1
        return 2 * numerator / (degree * (degree-1))
# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nice(G,n) for n in G)
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
def cc_nicer(G, v):
    {node: nx.clustering(G, node) if nx.clustering(G, node) != 0 else 0.0 for node in G.nodes()}

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nicer(G,n) for n in G)
```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    return nx.number_of_edges(G) == ((nx.number_of_nodes(G)) * (nx.number_of_nodes(G) - 1)) / 2
is_complete(G)
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```

```python

```