# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


```python
"""
vragen van First Course Chap 2
2.1 Heeft iedereen gedaan
2.2. True

2.3. True

2.4. Het minimum aantal links dat een ongericht netwerk zou kunnen hebben is N-1. Als het netwerk niet verplicht verbonden moet zijn, dan is het minimale aantal verbindingen 0.

2.5. True

2.6. False

2.7. True 

2.8. Het (undirected) netwerk in de adjacency matrix heeft een cycle (a - b - c - f - a) Het netwerk is niet sterk 
samenhangend omdat er geen pad is van e naar een andere knoop in het netwerk. Het netwerk is zwak samenhangend 
omdat we een pad kunnen vinden tussen elke twee nodes als we niet kijken naar de richting van de verbindingen. 
Als er een pad is tussen twee knooppunten in de ongerichte versie, is er altijd een pad tussen dezelfde twee knooppunten 
in de gerichte versie.

2.9 Nee, dit netwerk is geen tree. Het is namelijk na het weghalen van bepaalde relatie nog mogelijk om die punten te 
bereiken. Als bijvoorbeeld de relatie tussen D en F weggehaald wordt, kan er nog steeds van D naar F worden gegaan via E. 
Daarnaast heeft het netwerk meer links dan N – 1.
2.10. Diameter = 2
2.11. Ja,  als een directed network weakly connected is betekent het dat het netwerk wel met met links aan elkaar zit, 
maar dat je door de directed links niet overal kan komen. Oftewel het netwerk is connected als het undirected zou zijn. 
Er wordt gevraag of een weakly connected netwerk dat naar een undirect netwerk wordt omgezet dan connected is. 
Dit is dus het geval.
2.12 C
2.13 B
2.14 In-strength node D: 
    9 Out-strength node C: 3
2.15 1 (volgens nx)
2.16 B
2.17. 1, 2, 4, 6
2.18. 1, 2, 4, 5, 6
2.19. C
2.20. A
2.21. Nee, aangezien deze waarden duiden op een netwerk dat zowel verspreid als ‘sparse’ is qua verbindingen. 
Een sociaal netwerk is heeft juist vaak connecties tussen verschillende neighbours waardoor het minder ‘sparse’ 
zou moeten zijn.
2.22. A
2.23. 3
2.24. 2
2.25 - True
2.26 - D
2.27 - False
2.28 - A
2.29 - Chris Zaal - Karen I. Aardal = 3
    Paul Richard Halmos - 	Jan-Pieter L. Dorsman =  4
    Hessel B. Posthuma - Richard S. Hamilton = 4
2.30 - Average path length = 3
2.31 - 8 clicks 
2.32 - 1
2.33 In a tree, the maximum clustering coefficient for a node is 0. In a tree, each node has at most one parent, which means that it is impossible for a node to have two neighbors that are also neighbors of each other. In other words, there are no triangles in a tree, and therefore the clustering coefficient of any node in a tree is always 0.

2.34  The ego network of a node  is the subnetwork consisting of the chosen node called the ego- and its neighbors. Clustering coefficient = 0.3

2.35  A, B, C, D, E, F, G, H

A - B = 1
A - B - C = 2
A - B - D = 2
A - B - D - E = 3
A - B - D - F = 3
A - B - D - F - G = 4
A - B - D - F - H = 4

B - C = 1
B - D = 1
B - D - E = 2
B - D - F = 2
B - D - F - G = 3
B - D - F - H = 3

C - D - E = 2
C - D - F = 2
C - D - F - G = 3
C - D - F - H = 3

D - B = 1
D - C = 1
D - E = 1
D - F = 1
D - F - G = 2
D - F - H = 2

F - G = 1

2.36 Consider the undirected network in Figure 2.4. Compute the clustering coefficient for each node such that it is defined.

Cc of A = 0
Cc of B = 1/3
Cc of C = 1
Cc of D = 1/6
Cc of E = 0
Cc of F = 0
Cc of G = 0
Cc of H = 0

2.37 Shortest-path:
d - e = 1
d - e - c = 2
d - e - b = 2
d - e - c - g = 3
d - e - b - g = 3
d - e - b - h = 3
d - e - c/b - g - f = 4
d - e - c/b - g - a = 4

e - c = 1
e - b = 1
e - c - g = 2
e - b - g = 2
e - b - h = 2
e - c/b - g - f = 3
e - c/b - g - a = 3

c - b = 1
c - g = 1
c - b - h = 2
c - g - f = 2
c - g - a = 2

b - h = 1
b - g = 1
b - g - f = 2
b - g - a = 2

g - f = 1
g - a = 1
h - b - g = 2
h - b - g - a = 3
h - b - g - f = 3

Network average shortest-path length (L):
    n = 8
    d_total = 60 
    L = 1,071

2.38 Cc of a = 1
    Cc of b = ⅓
    Cc of c = ⅔ 
    Cc of d = 0
    Cc of e = ⅓ 
    Cc of f = 1
    Cc of g = ⅓
    Cc of h = 0

2.39 My friends: Sofie, Lotte, Sanne, Cherissa, Mike, Niffyn
    k  = 6
    Sofie neighbours: me, Sanne and Cherissa
    Lotte neighbors: me, Niffyn and Mike
    E = 4
    Cc =  4/15 

2.40
a. True
b. False 
c. False
d. True

"""
```

```python
# hieronder de vragen die gecheckt moetsen worden met networkx
FC = nx.DiGraph()
nodes = ['A', 'B', 'C', 'D', 'E', 'F']
edges = [('A', 'B'), ('B', 'C'), ('D', 'B'), ('D', 'E'), ('E', 'F'), ('F', 'A'), ('F', 'B'), ('F', 'C'), ('F', 'D'), ('F', 'E')]
FC.add_nodes_from(nodes)
FC.add_edges_from(edges)
#nx.draw(FC, with_labels = True)
# 2.08
nx.find_cycle(FC)
# geeft een cycle, dus de graph heeft een cycle
nx.is_weakly_connected(FC)
# nx.is_weakly_connected(FC) geeft true
nx.is_strongly_connected(FC)
# nx.is_strongly_connected(FC) geeft False
# de graph is due weakly connected
FC = FC.to_undirected(as_view=False)
nx.draw(FC, with_labels = True)
# 2.09
# een tree heeft N - 1 links
def is_tree(G) -> bool:
    """ returns true als het een tree is"""
    return G.number_of_edges() == G.number_of_nodes() - 1
print(nx.is_tree(FC))
# FC is doordat is_tree false geeft ook geen tree
# 2.10
# longest shortest path, nog uitzoeken
nx.diameter(FC)
# geeft 2 oftewel de diamter van FC is 2
# 2.11
WC = nx.DiGraph()
nodes = [1, 2, 3, 4]
edges = [(2,1), (1,3), (3,4), (4,1)]
WC.add_nodes_from(nodes)
WC.add_edges_from(edges)
#nx.draw(WC, with_labels = True)
nx.is_weakly_connected(WC)
# geeft True
WC = nx.to_undirected(WC)
nx.is_connected(WC)
# geeft ook true
# Doordat networkx aangeeft dat de directed weakly connected is en na het undirected maken van het netwerk geeft networkx
# dat het netwerk connected is betekent dit dat het eerder gegeven antwoord klopt.
```

```python
# Figuur 2.13
TD = nx.DiGraph()
nodes = ['A', 'B', 'C', 'D', 'E', 'F']
TD.add_nodes_from(nodes)
TD.add_edge("A", "B", weight=2)
TD.add_edge("A", "C", weight=6)
TD.add_edge("B", "F", weight=11)
TD.add_edge("B", "D", weight=9)
TD.add_edge("B", "C", weight=1)
TD.add_edge("C", "E", weight=3)
TD.add_edge("D", "F", weight=2)
TD.add_edge("E", "D", weight=2)
nx.draw(TD, with_labels = True)
# vraag 2.13
nx.is_weakly_connected(TD)
# geeft True, dus het gegeven antwoord klopt.
# vraag 2.14
# vraag 2.15
largest = max(nx.strongly_connected_components(TD), key=len)
print(len(largest))
```

```python
# figuur 2.14
TF = nx.DiGraph()
nodes = [1, 2, 3, 4, 5, 6]
TF.add_nodes_from(nodes)
TF.add_edge(1, 2, weight=0.8)
TF.add_edge(1, 3, weight=0.6)
TF.add_edge(2, 4, weight=0.9)
TF.add_edge(3, 5, weight=0.5)
TF.add_edge(4, 5, weight=0.8)
TF.add_edge(4, 6, weight=0.3)
TF.add_edge(5, 6, weight=0.95)
#nx.draw(TF, with_labels = True)
# vraag 2.16
nx.is_weakly_connected(TF)
# geeft treu, dus is figuur 2.14 weakly connected en daarmee is ons antwoord goed.
# 2.17
nx.shortest_path(TF, 1, 6)
# 2.18
# zelfde netwerk met inverse gewicht
TT = nx.DiGraph()
nodes = [1, 2, 3, 4, 5, 6]
TT.add_nodes_from(nodes)
TT.add_edge(1, 2, weight=(0.8 ** -1))
TT.add_edge(1, 3, weight=(0.6 ** -1))
TT.add_edge(2, 4, weight=(0.9 ** -1))
TT.add_edge(3, 5, weight=(0.5 ** -1))
TT.add_edge(4, 5, weight=(0.8 ** -1))
TT.add_edge(4, 6, weight=(0.3 ** -1))
TT.add_edge(5, 6, weight=(0.95 ** -1))
nx.draw(TT, with_labels = True)
nx.shortest_path(TT, 1, 6)
```

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    #WRITE YOUR CODE HERE
    return max(nx.strongly_connected_components(G), key = len)

def find_OUT(G):
    #WRITE YOUR CODE HERE
    empty_set = set()
    for nodes in find_SCC(G):
        for x in G:
            if x not in find_SCC(G):
                if nx.has_path(G, nodes, x):
                    empty_set.add(x)
    return empty_set

def find_IN(G):
    #WRITE YOUR CODE HERE
    empty_set = set()
    for nodes in find_SCC(G):
        for x in G:
            if x not in find_SCC(G):
                if nx.has_path(G, x, nodes):
                    empty_set.add(x)
    return empty_set
                
# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
#WRITE YOUR CODE HERE

def clusteringcoefficient(G):
    #WRITE YOUR CODE HERE
    empty_dict = dict()
    for nodes in G:
        empty_list = []
        for neighbors in G.neighbors(nodes):
            empty_list.append(neighbors)
        x = nx.subgraph(G, empty_list)
        if ((len(x) * (len(x) - 1)) / 2) != 0:
            empty_dict[nodes] = float(nx.number_of_edges(x) / ((len(x) * (len(x) - 1)) / 2))
        else: 
            empty_dict[nodes] = float(0)
        #nx.density(x)
    return empty_dict

# even er een mooi plotje van maken
pd.Series( clusteringcoefficient(G) ).sort_values().plot(kind='barh',figsize=(8,8))

clusteringcoefficient(G)
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
def cc(G, n):
    """ cc is de eerder gemaakte functie clusteringcoefficient, maar dan van 1 punt opgeroepen met key n"""
    return clusteringcoefficient(G)[n]

def cc_nice(G, v):
    edges = 0
    number = 0
    list_of_neighbors = []
    for nodes in G:
        if G.has_edge(v, nodes):
            list_of_neighbors.append(nodes)
    for nodes in list_of_neighbors:
        for nodes2 in list_of_neighbors:
            if nodes > nodes2:
                if G.has_edge(nodes, nodes2):
                    edges += 1
    if (len(list_of_neighbors) * (len(list_of_neighbors) - 1) / 2) == 0:
        return 0
    else:
        return float(edges / (len(list_of_neighbors) * (len(list_of_neighbors) - 1) / 2))
    
# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nice(G,n) for n in G)
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
def cc_nicer(G, v):
    neighbors_list = []
    for nodes in G:
        if (nodes, v) in G.edges():
            neighbors_list.append(nodes)
    sub_graph = G.subgraph(neighbors_list)
    return nx.density(sub_graph) 
# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nicer(G,n) for n in G)
```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    #WRITE YOUR CODE HERE
    return G.number_of_edges() == ((nx.number_of_nodes(G) * (nx.number_of_nodes(G) - 1)) / 2)
print(is_complete(G))
# test met een compleet netwerk
CN = nx.Graph()
nodes = [1, 2, 3, 4]
edges = [(1, 2), (1, 3), (1, 4), (2, 3), (2, 4), (3, 4)]
CN.add_nodes_from(nodes)
CN.add_edges_from(edges)
is_complete(CN)
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```

```python

```