# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


* 2.1 done
* 2.2 True
* 2.3 True
* 2.4 N-1
* 2.5 True
* 2.6 True
* 2.7 False
* 2.8
    * Dit netwerk heeft een cycle: (F,D) (D,E) (E,F)
    * Dit netwerk is weakly connected 
* 2.9 Dit netwerk is geen Tree
* 2.10 de diameter is 2
* 2.11 Ja, omdat er dan een pad is tussen elk punt
* 2.12 c
* 2.13 b
* 2.14 
    * D in-strenght = 11 
    * C out-stenght = 3
* 2.15 3
* 2.16 b
* 2.17 [1, 2, 4, 6]
* 2.18 [1, 2, 4, 5, 6]
* 2.19 c
* 2.20 a
* 2.21 nee, de clustering coefficient zou hoger zijn en de diameter kleiner.
* 2.22 a
* 2.23 6
* 2.24 2
* 2.25 nee, omdat elk pad dat bestaat in een directed network ook bestaat in een undirected network, dus de diameter en het kortste pad kunnen niet korter zijn.
* 2.26 c
* 2.27 True
* 2.28 a
* 2.29 M. Marx = 2
* 2.30 .
* 2.31 gemmidelde lengte = 4.25
* 2.32 1
* 2.33 0
* 2.34 0.3
* 2.35 
* 2.36
* 2.37
    * 1 ag/af/fg/gc/gb/cb/ce/bh/ed/eb
    * 2 ab/ac/fc/fb/ge/gh/ch/cd/bd/eh
    * 3 ae/ah/fe/fh/gd/dh
    * 4 ad/fd
    * gemiddelde average path length = 2
* 2.38
    * a 1
    * b 1/3
    * c 2/3
    * d 1
    * e 1/3
    * f 1
    * g 1/3
    * h 1
    * network: ongeveer 0.708
* 2.39 Ik heb een discordgroep met 6 van mijn vrineden gebruikt en mijn clustercoefficient is 0.9
* 2.40 d

```python
# Vragen 2.8-2.11
# Vraag 2.8
# Dit netwerk heeft een cycle: (F,D) (D,E) (E,F)
# Dit netwerk is weakly connected 
K = nx.DiGraph()

K.add_nodes_from(['A','B','C','D','E','F'])

K.add_weighted_edges_from([('A','B',1),('B','C',2),('D','B',1),('D','E',1),('E','F',1),('F','A',2),('F','B',1),('F','C',3),('F','D',1),('F','E',1)])

print('Het netwerk van vraag 2.8 heeft cycles')
print(nx.algorithms.dag.has_cycle(K))
print('Het netwerk van vraag 2.8 is strongly connected')
print(nx.is_strongly_connected(K))
print('Het netwerk van vraag 2.8 is weakly connected')
print(nx.is_weakly_connected(K))

# Vraag 2.9
# Dit netwerk is geen Tree
K_unweighted = nx.Graph()

K_unweighted.add_nodes_from(['A','B','C','D','E','F'])

K_unweighted.add_edges_from([('A','B'),('B','C'),('D','B'),('D','E'),('E','F'),('F','A'),('F','B'),('F','C'),('F','D'),('F','E')])

print('Het netwerk van vraag 2.9 is een tree')
print(nx.is_tree(K_unweighted))

# vraag 2.10
# de diameter is 2
print('De diameter van het netwerk van vraag 2.10 is:')
print(nx.diameter(K_unweighted))

#vraag 2.11
# Ja, omdat er dan een pad is tussen elk punt
print('Het netwerk van vraag 2.11 is connected')
print(nx.is_connected(K_unweighted))
```

```python
dertien = nx.DiGraph()
dertien.add_nodes_from(['A','B','F','D','E','C'])
dertien.add_weighted_edges_from([('A','B',2),('B','C',1),('B','D',9),('E','D',2),('B','F',11),('D','F',2),('C','E',3),('A','C',6)])

pos = nx.circular_layout(dertien)
edge_labels = nx.get_edge_attributes(dertien, 'weight')
nx.draw_networkx_edge_labels(dertien, pos, edge_labels=edge_labels, label_pos=0.5)
nx.draw(dertien, pos, with_labels=True)
```

```python
veertien = nx.DiGraph()
veertien.add_nodes_from([6,4,2,1,3,5])
veertien.add_weighted_edges_from([(1,2,0.8),(1,3,0.6),(2,4,0.9),(3,5,0.5),(4,5,0.8),(4,6,0.3),(5,6,0.95)])

pos = nx.circular_layout(veertien)
edge_labels = nx.get_edge_attributes(veertien, 'weight')
nx.draw_networkx_edge_labels(veertien, pos, edge_labels=edge_labels, label_pos=0.5)
nx.draw(veertien, pos, with_labels=True)
```

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    return max([i for i in nx.strongly_connected_components(G)], key=len)

def find_OUT(G):
    return {i for i in G for j in find_SCC(G) if G.has_edge(j, i) and i not in find_SCC(G)}
    
    
def find_IN(G):
    return {i for i in G for j in find_SCC(G) if G.has_edge(i, j) and i not in find_SCC(G)}
    
# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

```python

```

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
def clusteringcoefficient(G):
    
    return {x: ((len((G.subgraph([y for y in G.neighbors(x)])).edges()) / 
                (len(list(G.neighbors(x))) * ((len(list(G.neighbors(x))) - 1) / 2))) 
                if len(list(G.neighbors(x))) > 1 else float(0)) 
            for x in G}

# even er een mooi plotje van maken
pd.Series( clusteringcoefficient(G) ).sort_values().plot(kind='barh',figsize=(8,8))

clusteringcoefficient(G)
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

```python

```

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
def cc_nice(G, v):
    neighbors = []
    for i in G:
        if G.has_edge(i,v):
            neighbors.append(i)
    
    triangles = 0
    for i in neigbors:
        for j in neighbors:
            if G.has_edge(i,j):
                triangles += 1
    
    return triangles / (((len(G)) * (len(G) - 1)) / 2 )

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
#all(cc(G,n)==cc_nice(G,n) for n in G)
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
WRITE YOUR CODE HERE

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nicer(G,n) for n in G)
```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    if G.number_of_edges() == ((nx.number_of_nodes(G)) * (nx.number_of_nodes(G) - 1)) / 2:
        return True
    else:
        return False
is_complete(G)
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```

```python

```