# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


```python
"""
2.2 True
2.3 True
2.4 One link. Not connected -> zero links
2.5 False
2.6 False
2.7 True
2.8 EF, DF are cycles. It is strongly connected
2.9 It is not a tree.
2.10 4
2.11 Yes, you can only convert a connected network to an unconnected
network by adding nodes or removing links.
2.12 C
2.13 A
2.14 In-strength of D: 11, Out-strength of C: 3
2.15 - 3
2.16 - B
2.17 - The shortest path from node 1 to node 6 has a total length of 2 
2.18 - the shortest path (the reciprocals of the link weights) between 
nodes 1 and 6 has a total distance of 3.45 
2.19 - B
2.20 - A
2.21 - A social network is unlikely to have the diameter and clustering 
coefficient of the graph in Figure 2.15. Social networks have a higher 
clustering coefficient, indicating that the network's nodes are more tightly 
connected. Furthermore, because people are generally only a few degrees of 
separation apart, social networks have a smaller diameter.
2.22 - A
2.23 - Diameter is 4
2.24 -  Diameter is undefined since you can go back and forth along the same edge
2.25 - False, the directed graph D is a subset of the undirected graph G, 
with two undirected edges going in opposite directions replacing each directed edge. 
Because the directed graph D contains all of the edges of the undirected graph G, 
its shortest-path lengths and diameter cannot be less than G's.
2.26 - C
2.27 - True, the diameter of a network is defined as the shortest path between 
any two nodes on the network. As a result, the diameter is always greater than 
or equal to the average path length, which is the sum of the network's shortest paths.
2.28 D
2.29 Raf Bocklandt - Paul Erdős: Erdos number = 4
       Maarten Marx - Balder ten Cate: Collaboration Distance = 1
2.30 Makoto Shinkai - Jackie Chan: Distance = 4
       Kana Hanazawa - Chris Tucker: Distance = 3
       Seth MacFarlane - AngelaBaby: Distance = 2
       Christian Bale - Eva Green: Distance = 2
       Rie Takahashi - Felix Roco: Distance = 4
       Timothy Olyphant - Qorianka Kilcher: Distance = 2
       Yuki Kaji - Michael B. Jordan: Distance = 3
       IU - Richard Pepple: Distance = 3
       Daisuke Ono - Park Sooyoung: Distance = 4
       Lewis Hamilton - Nanoka Hara: Distance = 3
       Average path length: 3.0
       Modus: 3

2.31 Ronde 1: Warship - Airport: Length of Path: 5
       Ronde 2: Peter Sellers - European Commission: 6
       Ronde 3: Internet Meme - Computer animation: 6
       Ronde 4: Pig - John F. Kennedy International Airpot: 5
       Ronde 5: Steve Jobs - Anesthesia: 7
       Average Path Length: 5.8
        
2.32 Max. is 1

2.33 0

2.34 C = (2 * 3) / 5(5 - 1) = 6 / 20 = 0.3

2.35 Meest linker Node: 1, 2, 2, 3, 3, 4, 4
       Direct rechtsliggend: 1, 1, 1, 2, 2, 3, 3
       Op schuine zijde liggende note: 1, 1, 2, 2, 2, 3, 3
       Middelste node: 1, 1, 1, 1, 2, 2, 2
       Onder middelste node: 1, 2, 2, 2, 3, 3, 3
       Rechts van middelste node: 1, 1, 1, 2, 2, 2, 3
       Hieronder: 1, 2, 2, 3, 3, 3, 4
       Helemaal rechts: 1, 2, 2, 3, 3, 3, 4
        
2.36 C = (2 * 1) / (2(2 - 1)) = 1
       C = (2 * 1) / (3(3 - 1)) = 1 / 3
       C = (2 * 1) / (4(4 - 1)) = 1 / 6
       C = 0 
    
2.37 a,g = 1; a,f= 1; a,c = 2; a,b = 2; a,h = 3; a,e = 3; a,d = 4
       b,e = 1; b,h =1; b,g = 1; b,c = 1; b,d = 2; b,f = 2
       c,g = 1; c,e = 1; c,d = 2; c,h = 2; c,f = 2
       d,e = 1; d,h = 3; d,g = 3; d,f = 4
       e,g = 2; e,h =2; e,f = 3
       f,g = 1; f,h = 3
       g,h = 2
       Average = 56 / 28 = 2
    
2.38 C(a) = (2 * 1) / (2(2 - 1)) = 1
       C(b) = (2 * 2) / (4(4 - 1)) = 1/3
       C(c) = (2 * 2) / (3(3 - 1)) = 2/3
       C(e) = (2 * 1) / (3(3 - 1)) = 1/3
       C(f) = (2 * 1) / (2(2 - 1)) = 1
       C(g) = (2 * 2) / (4(4 - 1)) = 1/3
       Average(C) = 0.611
        
2.39 Ik heb geen vrienden lmao

2.40 D
"""
```

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    sccs = list(nx.strongly_connected_components(G))
    max_scc = max(sccs, key=len)
    return set(max_scc)

scc = find_SCC(EK)
in_nodes = set()

def dfs(node):
    in_nodes.add(node)
    for neighbor in EK.predecessors(node):
        if neighbor not in in_nodes:
            dfs(neighbor)

    out_nodes = set()
    for node in scc:
        path_nodes = nx.descendants(EK.reverse(), node)
        path_nodes.add(node)
        path = EK.subgraph(path_nodes)
        if not nx.find_cycle(path):
            out_nodes |= set(nx.topological_sort(path))

    for node in scc:
        dfs(node)

def find_SCC(G):
    return set(max(nx.strongly_connected_components(G), key=len))

def find_OUT(G):
    out_nodes = set()
    return out_nodes - scc

def find_IN(G):
    return in_nodes - scc
    
# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
def clusteringcoefficient(G):
    cc = {} 
    
    for node in G.nodes():
        neighbors = list(G.neighbors(node))
        k_i = len(neighbors)
        
        # no clustering coefficients for nodes with k_i < 2
        if k_i < 2:  
            cc[node] = 0
        else:
            t_i = 0
            for i, neighbor1 in enumerate(neighbors):
                for neighbor2 in neighbors[i+1:]:
                    if G.has_edge(neighbor1, neighbor2):
                        t_i += 1
            
            cc[node] = float((2 * t_i) / (k_i * (k_i - 1)))
    
    return cc

G = nx.gnm_random_graph(20, 50)

# compute the clustering coefficients for each node
cc = clusteringcoefficient(G)

# print the clustering coefficients
for node in G.nodes():
    print(f"Node {node}: {cc[node]}")

# even er een mooi plotje van maken
pd.Series( clusteringcoefficient(G) ).sort_values().plot(kind='barh',figsize=(8,8))

clusteringcoefficient(G)
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
WRITE YOUR CODE HERE

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nice(G,n) for n in G)
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
WRITE YOUR CODE HERE

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nicer(G,n) for n in G)
```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    n = nx.number_of_nodes(G)
    m = G.number_of_edges()
    return m == n*(n-1)//2
is_complete(G)
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```

```python

```