# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


### Antwoorden
2.1: Gedaan<br>
2.2: Als het netwerk verbonden is, is er altijd een kortste pad tussen twee verschillende nodes.<br>
2.3: Ja, dat klopt, dat pad gaat altijd door de root van de tree. Daarnaast is de lengte van het pad dus ook de lengte van node A tot de root + de lengte van node B tot de root.<br>
2.4: Het minimum aantal links dat een connected network kan hebben is (N-1), wanneer het netwerk niet connected <br>
2.5: Ja, want een minimaal connected netwerk heeft ook (N-1) en hier kunnen ook geen cycles in plaatsvinden, dus moet het wel een tree zijn.<br>
2.6: Ja, dit is het geval, want in een netwerk met minimaal N links moet er minimaal een node zijn die 2 links heeft. Dit zorgt dus voor dat er wel een cycle moet zijn.<br>
2.7: Nee dit hoeft niet.<br>
2.8: Ja er is een cycle, bijvoorbeeld van D->E, E->F, F->D. Hij is weakly connected, maar niet strongly connected.<br>
2.9: Nee, dit netwerk is geen tree, omdat er meerdere cycles in voorkomen.<br>
2.10: De diameter van het netwerk is 2.<br>
2.11: Ja, want weakly connected betekent dat de direction van links niet uitmaakt. Dit betekent dus dat als het netwerk undirected wordt gemaakt dat hij alsnog connected is.<br>
2.12: C<br>
2.13: B<br>
2.14: In-strenght node D is 11, out-strength node C is 3.<br>
2.15: 1 node.<br>
2.16: B<br>
2.17: 1->2->4->6<br>
2.18: 1->2->4->5->6<br>
2.19: C<br>
2.20: A<br>
2.21: Nee, want een sociaal netwerk heeft waarschijnlijk veel meer gesloten driehoeken door de strong triadic closure property. Hierdoor is de diameter van het netwerk waarschijnlijk lager en wordt de clusteringcoëfficent groter.<br>
2.22: A<br>
2.23: Een diameter van 5<br>
2.24: Een diameter van 2<br>
2.25: Nee, want in een directed netwerk met N=2 is de average path lengte en diameter even groot als in een undirected netwerk met N=2<br>
2.26: C<br>
2.27: Ja dat klopt, want de diameter is het langste kortste pad, dus dit kan nooit minder zijn dan het gemiddelde padlengte.<br>
2.28: A<br>
2.29: Gedaan<br>
2.30: Gedaan<br>
2.31: Gedaan<br>
2.32: 1<br>
2.33: 0<br>
2.34: 2x3/5(5-1) = 6/20 = 0.3<br>
2.35: Gedaan<br>
2.36: Gedaan<br>
2.37: Gedaan<br>
2.38: Gedaan<br>
2.39: -<br>
2.40: D<br>

```python
G8_11 = nx.DiGraph()
G8_11.add_weighted_edges_from([('A','B',1), ('B','C',2), ('D','B',1), ('D','E',1), ('E','F',1), ('F','A',2), ('F','B',1), ('F','C',3), ('F','D',1), ('F','E',1)])
nx.draw(G8_11, with_labels=True)

#2.8
print(nx.is_tree(G8_11)) #False, dus cycles
print(nx.is_strongly_connected(G8_11)) #False
print(nx.is_weakly_connected(G8_11)) #True

```

```python
G8_11_unweighted_undirected = nx.Graph()
G8_11_unweighted_undirected.add_edges_from([("A","B"), ("A","F"), ("B","C"), ("B","D"), ("B","F"), ("C","F"), ("D","E"), ("D","F"), ("E","F")])
nx.draw(G8_11_unweighted_undirected, with_labels=True)

#2.9
print(nx.is_tree(G8_11_unweighted_undirected))

#2.10
print(nx.diameter(G8_11_unweighted_undirected))

#2.11
print(nx.is_connected(G8_11_unweighted_undirected)) 

```

```python
G2_13 = nx.DiGraph()
G2_13.add_weighted_edges_from([('A','B',2), ('B','F',11), ('D','F',2), ('E','D',2), ('C','E',3), ('A','C',6), ('B','C',1), ('B','D',9)])

pos=nx.circular_layout(G2_13) # pos = nx.nx_agraph.graphviz_layout(G)
nx.draw_networkx(G2_13,pos)
labels = nx.get_edge_attributes(G2_13,'weight')
nx.draw_networkx_edge_labels(G2_13,pos,edge_labels=labels)
```

```python
G2_14 = nx.DiGraph()
G2_14.add_weighted_edges_from([('1','2',0.8), ('2','4',0.9), ('4','6',0.3), ('1','3',0.6), ('3','5',0.5), ('5','6',0.95), ('4','5',0.8)])

pos=nx.kamada_kawai_layout(G2_14)
nx.draw_networkx(G2_14,pos)
labels = nx.get_edge_attributes(G2_14,'weight')
nx.draw_networkx_edge_labels(G2_14,pos,edge_labels=labels)
```

```python
#2.14
print(G2_13.in_degree('D', weight='weight'))
print(G2_13.out_degree('C', weight='weight'))

#2.15
len(max(nx.strongly_connected_components(G2_13), key=len))

#2.17
nx.shortest_path(G2_14, '1', '6', weight='weight')

```

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    return max(nx.strongly_connected_components(EK), key=len)

def find_OUT(G):
    #WRITE YOUR CODE HERE
    return {node for node in EK if (nx.has_path(EK, list(find_SCC(EK))[0], node)) and not nx.has_path(EK, node, list(find_SCC(EK))[0])}

def find_IN(G):
    #WRITE YOUR CODE HERE
    return {node for node in EK if (nx.has_path(EK, node, list(find_SCC(EK))[0])) and not nx.has_path(EK, list(find_SCC(EK))[0], node)}

# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
def clusteringcoefficient(G):    
    clustering_dict = {}
    for node in G:
        if G.degree(node) == 1:
            clustering_dict[node] = 0.0
            continue
        triangles = {edge for edge in G.edges() if edge[0] in G.neighbors(node) and edge[1] in G.neighbors(node)}
        clustering_dict[node] = (2*len(triangles))/(G.degree(node)*(G.degree(node)-1))
    return clustering_dict

# even er een mooi plotje van maken
pd.Series( clusteringcoefficient(G) ).sort_values().plot(kind='barh',figsize=(8,8))

clusteringcoefficient(G)
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
def cc_nice(G, n):
    neighbor_set = {node for node in G if G.has_edge(node, n)}
    
    if len(neighbor_set) <= 1:
        return 0
    
    triangles = {frozenset({node1, node2}) for node1 in neighbor_set for node2 in neighbor_set if G.has_edge(node1, node2)}
    return (2*len(triangles))/(len(neighbor_set)*(len(neighbor_set)-1))
    
#check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(nx.clustering(G,n)==cc_nice(G,n) for n in G)
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
def cc_nicer(G, n):
    return nx.density(nx.subgraph(G, {node for node in G if node in G[n]}))


# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(nx.clustering(G,n)==cc_nicer(G,n) for n in G)
```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

```python
H = nx.complete_graph(6)
nx.draw(H)
```

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    return G.number_of_edges() == ((G.number_of_nodes() * (G.number_of_nodes() - 1))/2)

print(is_complete(G))
print(is_complete(H))
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```