# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


1 check,

2 False, want er bestaan tussen 2 nodes maar maximaal 1 edge

3 True
 
4 Aantal nodes - 1, want dan kan je vanuit iedere node naar een andere node komen = verbonden.

5 True

6 False

7 True

8 ?

9 ?

10 ?

11 Ja, want dan kan je vanuit iedere node naar een andere node komen ongeacht de richting. 

12 C 

13 C

14 In strenght = 11 Out strength = 3

15 2 

16 C
 
17 1-2-4-6

18 ?

19 C

20 B

21 Yes

22 A

23 3

24 2
 
25 False

26 C

27 True

28 D

29 Check

30 Check

31 3

32 1 

33 ?

34 0.2

35 prog

36 prog

37 prog

38 prog

39 prog

40 prog

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    return set(max(nx.strongly_connected_components(G), key=len))

def find_OUT(G):
    SCC = find_SCC(G)
    out = set()
    for node in SCC:
        edges = nx.dfs_edges(G, node, 1)
        for edge in edges:
            if edge[1] not in SCC:
                out.add(edge[1])
    return out
        

def find_IN(G):
    SCC = find_SCC(G)
    out = set()
    for node in G.nodes:
        edges = nx.dfs_edges(G, node, 1)
        for edge in edges:
            if edge[1] in SCC and edge[0] not in SCC:
                out.add(edge[0])
    return out

# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

```python
def find_SCC(G):
    return set(max(nx.strongly_connected_components(G), key=len))

def find_OUT(G):
    SCC = find_SCC(G)
    out = set()
    niet = set(G.nodes()) - SCC
    for node1 in SCC:
        for node2 in niet:
            if nx.has_path(G, node1, node2):
                out.add(node2)
    return out
        

def find_IN(G):
    SCC = find_SCC(G)
    out = set()
    niet = set(G.nodes()) - SCC
    for node1 in SCC:
        for node2 in niet:
            if nx.has_path(G, node2, node1):
                out.add(node2)
    return out

# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

The “tendrils” of the bow-tie consist of (a) the nodes reachable from IN that
cannot reach the giant SCC, and (b) the nodes that can reach OUT but cannot be
reached from the giant SCC


Disconnected: Finally, there are nodes that would not have a path to the giant SCC
even if we completely ignored the directions of the edges. These belong to none of the
preceding categories.

 It’s possible for a tendril node to satisfy both (a) and (b), in which case it’s
part of a “tube” that travels from IN to OUT without touching the giant SCC.

```python
?nx.has_path
```

```python
?nx.dfs_edges
```

The nx.dfs_edges function is a method provided by the Python library NetworkX, which is used to perform a depth-first search (DFS) on a graph and return a generator of edges visited during the search.

The function takes a graph object G and a starting node source as input, and returns a generator that yields the edges of the graph in the order in which they are visited during the depth-first search. The function uses the DFS algorithm to traverse the graph and visit all nodes in the connected component of the source node.

```python
list(nx.strongly_connected_components(EK))
```

```python
?nx.in_degree_centrality
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

set.add(toevoegen)
list.append(toevoegen)
dict[key] = value

De teller is het aantal vriendschappen tussen de vrienden van n.

De noemer is het aantal mogelijke vriendschappen tussen de vrienden van n.

```python
?G.has_edge
```

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
def clusteringcoefficient(G):
    resultaat = dict()
    for node in G.nodes():
        if G.degree(node) == 1:
            resultaat[node] = 0
        else:
            edgesburen = 0
            for node1 in G.neighbors(node):
                for node2 in G.neighbors(node):
                    if node1 != node2:
                        if G.has_edge(node1, node2):
                            edgesburen += 1
                            
            degree = G.degree(node)
            teller = edgesburen
            noemer = degree * (degree - 1)
            formule = teller/noemer
            resultaat[node] = formule
            
    return resultaat 
# even er een mooi plotje van maken
pd.Series( clusteringcoefficient(G) ).sort_values().plot(kind='barh',figsize=(8,8))
clusteringcoefficient(G)
```

```python
print(nx.clustering(G))
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
def cc_nice(G,v):
    resultaat = dict()
    if G.degree(v) == 1:
        resultaat[v] = 0
    else:
        edgesburen = 0
        buren = set()
        for knoop in G:
            if v != knoop:
                if G.has_edge(v, knoop):
                    buren.add(knoop)

        for node1 in buren:
            for node2 in buren:
                if node1 != node2:
                    if G.has_edge(node1, node2):
                        edgesburen += 1

        degree = G.degree(v)
        teller = edgesburen
        noemer = degree * (degree - 1)
        formule = teller/noemer
        resultaat[v] = formule

    return resultaat 

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nice(G,n) for n in G)
```

```python
?nx.triangles
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
def cc_nicer(G,v):
    resultaat = dict()
    if G.degree(v) == 1:
        resultaat[v] = 0
    else:
        edgesburen = 0
        buren = set()
        for knoop in G:
            if v != knoop:
                if G.has_edge(v, knoop):
                    buren.add(knoop)

        for node1 in buren:
            for node2 in buren:
                for node3 in buren:
                    if node1 != node2:
                        if node1 != node3:
                            if nx.density(G, node1, node2, node3) == 1:
                                edgesburen += 1

        degree = G.degree(v)
        teller = edgesburen
        noemer = degree * (degree - 1)
        formule = teller/noemer
        resultaat[v] = formule

    return resultaat 

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
cc_nice(G,v)
```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    N = nx.number_of_nodes(G)
    E = G.number_of_edges()
    teller = N * (N - 1)
    noemer = 2
    maxedges = teller / noemer
    print(maxedges)
    print(E)
    if maxedges == E:
        return True
    else:
        return False
    
is_complete(G)
```

```python
def is_complete(G):
    N = nx.number_of_nodes(G)
    E = G.number_of_edges()
    teller = N * (N - 1)
    noemer = 2
    maxedges = teller / noemer
    
    density = (2 * E) / teller
    print(density)
    
    if density == 1:
        return True
    else:
        return False
is_complete(G)
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```

```python

```