# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


2.2 Recall that unless otherwise specified, the length of a path is the number of links contained therein. Given two nodes in an arbitrary undirected, connected graph, there must exist some shortest path between them. True or False: There may exist multiple such shortest paths.

    True
    
2.3 True or False: Given any two nodes in an (undirected) tree, there exists exactly one path between those two nodes.

    True
    
2.4 Consider an undirected, connected network with N nodes. What is the minimum number of links the network can have? If we do not require the network to be connected, does that minimum number of links change?

    Het minimum aantal links die het netwerk kan hebben is N - 1. Als het netwerk niet verbonden hoeft te zijn, betekent dit dat niet elke knoop via een of ander pad bij elke andere knoop kan komen. Dit betekent dat er paden, en daarmee links, kunnen missen. Hierdoor verandert het minimum aantal links.
    
2.5 Recall that a tree of N nodes contains N — 1 links. True or False: Any connected, undirected network of N nodes and N — 1 links must be a tree.

    False
    
2.6 True or False: Any undirected network of N nodes with at least N links must contain a cycle.

    True
    
2.7 True or False: Any directed network of N nodes with at least N links must contain a cycle.

    False

2.8 Consider the network defined by the adjacency matrix in Eq. (1.11). Are there any cycles in this network? Is it strongly connected? Weakly connected?

```python
D = nx.DiGraph()
D.add_weighted_edges_from([('A', 'B', 1), 
              ('B', 'C', 2),
              ('D', 'B', 1),
              ('D', 'E', 1),
              ('E', 'F', 1),
              ('F', 'A', 2),
              ('F', 'B', 1),
              ('F', 'C', 3),
              ('F', 'D', 1),
              ('F', 'E', 1)])

nx.draw(D, with_labels=True)

any_cycles = list(nx.find_cycle(D)) # [('D', 'E'), ('E', 'F'), ('F', 'D')]
# Ja, er bevindt zich een cycle in dit netwerk

nx.is_strongly_connected(D) # False
# Nee, het netwerk is niet strongly connected

nx.is_weakly_connected(D) # True
# Ja, het netwerk is weakly connected
```

2.9 Consider the unweighted, undirected version of the network defined by the adjacency matrix in Eq. (1.11). Is this network a tree?

```python
G = nx.to_undirected(D)

for node1, node2, weight in G.edges(data=True):
    del weight

nx.draw(G, with_labels=True)

nx.is_tree(G) # False
# Nee, het netwerk is geen tree
```

2.10 Consider the unweighted, undirected version of the network defined by the adjacency matrix in Eq. (1.11). What is this network’s diameter?

```python
nx.diameter(G)
```

2.11 If you convert a weakly connected directed network to an undirected network, will the resulting network be connected? Explain why or why not.

    Een weakly connected directed network bevat gerichte pijlen, maar worden verwaarloosd waardoor het netwerk verbonden wordt. Dit kan dus worden gezien als een ongericht netwerk. Aangezien het al eerder weakly verbonden was, is het netwerk als een ongericht netwerk ook verbonden.
    
2.12 Consider an arbitrary non-complete undirected network. Now add a single link. How has the number of nodes in this network’s giant component changed as a result of this addition?
a. It has strictly decreased
b. It has decreased or stayed the same
c. It has increased or stayed the same
d. It has strictly increased

    c

2.13 Consider the weighted directed network in Figure 2.13. Which of the following most accurately describes the connectedness of this network?
a. Strongly connected
b. Weakly connected
c. Disconnected
d. None of the above

    b

```python
D = nx.DiGraph()
D.add_weighted_edges_from([('A', 'C', 6), ('A', 'B', 2), ('B', 'C', 1), ('C', 'E', 3), ('B', 'D', 9), ('E', 'D', 2), ('B', 'F', 11), ('D', 'F', 2)])

pos = nx.spring_layout(D)
nx.draw(G, pos, with_labels=True)
my_edge_labels = {(u, v): d for u, v, d in G.edges(data=True)}

nx.draw_networkx_edge_labels(G, pos, edge_labels=my_edge_labels)
```

2.14 Consider the weighted directed network in Figure 2.13. What is the in-strength of node D? What is the out-strength of node C? (Recall the definitions from Chapter 1.)

```python
D.in_degree('D', weight='weight') # 11
D.out_degree('C', weight='weight') # 3
```

2.15 How many nodes are in the largest strongly connected component of the network in Figure 2.13?

```python
max(nx.strongly_connected_components(D), key=len)
```

2.16 Consider the network in Figure 2.14. Which of the following most accurately describes the connectedness of this network?
a. Strongly connected
b. Weakly connected
c. Disconnected
d. None of the above

    b

2.17 Link weights can represent anything about the relationship between the nodes: strength of the relationship, geographic distance, voltage flowing through a link cable, etc. When discussing path lengths on a weighted graph, one must first define how the weights are related to the distances. The length of a path between two nodes is then the sum of the distances of the links in that path. The simplest case occurs when the link weights represent the distance. Consider the network in Figure 2.14 and assume that the link weights represent distances. Using this distance metric, what is the shortest path between nodes 1 and 6?

```python
D = nx.DiGraph()
D.add_weighted_edges_from([(1, 2, 0.8), (1, 3, 0.6), (2, 4, 0.9), (3, 5, 0.5), (4, 5, 0.8), (4, 6, 0.3), (5, 6, 0.95)])

nx.draw(D, with_labels=True)

nx.shortest_path(D, 1, 6, 'weight')
```

2.18 A common way to define the distance between two nodes is the inverse (or reciprocal) of the link weight. Consider the network in Figure 2.14, and assume that the distance between two adjacent nodes is defined as the reciprocal of the link weight. Using this distance metric, what is the shortest path between nodes 1 and 6?

```python
for node1, node2, data in D.edges(data=True):
    data['weight'] = 1 / data['weight']

nx.shortest_path(D, 1, 6, 'weight')
```

2.19 Consider the network in Figure 2.15. Which of the following is the best estimate of this network’s diameter?
a. 2
b. 4
c. 10
d. 20
    
    c

2.20 Consider the network in Figure 2.15. Which of the following is the best estimate for the average clustering coefficient of this graph?
a. 0.05
b. 0.5
c. 0.75
d. 0.95

    a

2.21 Would a social network be likely to have the diameter and clustering coefficient of the graph in Figure 2.15?

    Nee

2.22 Consider the network in Figure 2.16. Which of the following most accurately describes the connectedness of this network?
a. Strongly connected
b. Weakly connected

    a

2.23 What is the diameter of the network in Figure 2.16?

```python
D = nx.DiGraph()
D.add_edges_from([(1, 2), (1, 4), (2, 4), (2, 5), (3, 1), (4, 3), (4, 6), (4, 7), (4, 5), (5, 7), (6, 3), (7, 6)])

nx.diameter(D)
```

2.24 Consider an undirected version of the network in Figure 2.16. What is the diameter of this network?

```python
G = nx.Graph(D)

nx.diameter(G)
```

2.25 Consider any arbitrary directed graph D along with its undirected version G. True or False: If the average shortest-path length and diameter of the directed graph exist, they can be smaller than those of the undirected version.

```python
print(nx.average_shortest_path_length(D))
print(nx.average_shortest_path_length(G))
```

    False

2.26 Imagine that you were building a competitor of NetworkX. You have already written a method shortest_path() to compute the shortest path between two nodes, and now you want to write a function to compute the diameter of a network. Which of the following best describes how to go about doing this?
a. First compute the shortest-path lengths between each pair of nodes. The diameter is the minimum of these values
b. First compute the shortest-path lengths between each pair of nodes. The diameter is the average of these values
c. First compute the shortest-path lengths between each pair of nodes. The diameter is the maximum of these values
d. First compute the average length of all paths between each pair of nodes. The diameter is the minimum of these values

    c

2.27 True or False: A network’s diameter is always greater than or equal to its average path length.

    False

2.28 What is the central idea behind the notion of “six degrees of separation’?
a. Social networks have high clustering coefficients
b. Social networks are sparse
c. Social networks have many high-degree nodes
d. Social networks have small average path length

    a

2.29 The American Mathematical Society has a Web tool to find the collaboration distance between two mathematicians (see Box 2.3). Use this tool to calculate the Erdös number for a few mathematicians in your institution, or whom you know by fame.
    
    Maarten Marx: 2
    Balder ten Cate: 3

2.30 Use The Oracle of Bacon (oracleofbacon.org) to measure the shortest-path distance in the co-star network among as many pairs of obscure actors and actresses as you can think of. Plot a histogram showing the distribution of the shortest-path lengths, and also estimate the average path length based on your sample. (If you are not familiar with histograms, they are defined in the next chapter.)

    Leonardo DiCaprio: 2
    Tom Cruise: 1
    Johnny Depp: 1
    Will Smith: 2
    Tom Hanks: 1
    Tom Holland: 2
    Zendaya: 2
    Julia Roberts: 1
    Angelina Jolie: 2
    Brad Pitt: 1
    Ian Somerhalder: 2
    Paul Wesley: 2
    Nina Dobrev: 2
    Ariana Grande: 2

```python
shortest_path_lenghts = [2, 1, 1, 2, 1, 2, 2, 1, 2, 1, 2, 2, 2, 2]

plt.hist(shortest_path_lenghts)
plt.xlabel('shortest path lengths')
plt.ylabel('frequency')
plt.show()

average_path_length = sum(shortest_path_lenghts) / (len(shortest_path_lenghts) * (len(shortest_path_lenghts) - 1))
print(average_path_length)
```

2.31 Play The Wiki Game (thewikigame.com) until you are able to complete a few rounds successfully. Report the average length (number of clicks) of the discovered paths.

    5

2.32 What is the maximum clustering coefficient for a node in an arbitrary undirected graph?

    1
    
2.33 What is the maximum clustering coefficient for a node ina tree?

    0
    
2.34 Recall the definition of an ego network in Section 1.4. Consider the ego network in Figure 2.17: what is the clustering coefficient of the ego?

    0,5
    
2.35 Consider the undirected network in Figure 2.4. Compute the shortest-path length for each pair of nodes in the giant component.

```python
G = nx.Graph()
G.add_edges_from([(1, 2), (2, 3), (2, 4), (3, 4), (4, 5), (4, 6), (4, 7), (4, 8)])
dict(nx.shortest_path_length(G))
```

2.36 Consider the undirected network in Figure 2.4. Compute the clustering coefficient for each node such that it is defined.

```python
nx.clustering(G)
```

2.37 Consider the network example in Figure 2.12. Compute the shortest-path length for each pair of nodes, and the average shortest-path length for the network.

```python
G = nx.Graph()
G.add_edges_from([('d', 'e'), ('e', 'c'), ('e', 'b'), ('c', 'b'), ('c', 'g'), ('b', 'h'), ('b', 'g'), ('g', 'f'), ('g', 'a'), ('f', 'a')])

print(dict(nx.shortest_path_length(G)))
nx.average_shortest_path_length(G)
```

2.38 Consider the network example in Figure 2.12. Compute the clustering coefficient for each node such that it is defined, and for the network.

```python
nx.clustering(G)
```

2.39 If you use an online social network such as Facebook or LinkedIn, measure your clustering coefficient in the network. (Hint: If you use a social network with directed links, such as Twitter or Instagram, you can treat the links as undirected.) (Hint 2: This might take a while; it’s okay to make an estimate based on a small sample of your friends.)
    
2.40 Which of the following seemingly conflicting properties are true of social networks?
a. Social networks have short paths, yet large diameter
b. Social networks have small diameter, yet large average path length
c. Social networks have many high-degree nodes, yet are disconnected
d. Social networks are highly clustered, yet are not dense

    a

2.41

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    "Geeft de set van nodes terug die tot de grootste SCC behoren"
    scc = max([c for c in nx.strongly_connected_components(G)], key=len)
    
    return set(scc)
    
def find_OUT(G):
    "Geeft de set met alle nodes die behoren tot de OUT"
    scc = find_SCC(G)
    list_OUT = []
    for node in G:
        if node not in scc:
            for s in scc:
                if G.has_edge(s, node) and not G.has_edge(node, s):
                    list_OUT.append(node)           
    return set(list_OUT)
    
def find_IN(G):
    "Geeft de set met alle nodes die behoren tot de IN"
    scc = find_SCC(G)
    list_IN = []
    for node in G:
        if node not in scc:
            for s in scc:
                if G.has_edge(node, s) and not G.has_edge(s, node):
                    list_IN.append(node)
    return set(list_IN)
            
    
# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
def clusteringcoefficient(G):
    "Geeft een dict die per knoop de clustering coefficient aangeeft"
    cc_dict = {}
    for node in G:
        degree = len(list(G.neighbors(node)))
        if degree >= 2:
            links = 0
            for i in G.neighbors(node):
                for j in G.neighbors(node):
                    if G.has_edge(i, j):
                        links += 1     
            cc_dict[node] = 2 * (links / 2) / (degree * (degree - 1))           
        else:
            cc_dict[node] = 0
            
    return cc_dict
                        
        
# even er een mooi plotje van maken
pd.Series(clusteringcoefficient(G)).sort_values().plot(kind='barh',figsize=(8,8))

clusteringcoefficient(G)
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
def cc_nice(G, v):
    "Geeft een dict die per knoop de clustering coefficient aangeeft"
    cc_dict = {}
    for node in G:
        list_neighbors = []
        degree = 0 
        for s in G:
            if node != s and G.has_edge(node, s):
                list_neighbors.append(s)
                degree += 1
                
        if degree >= 2:
            links = 0 
            for i in list_neighbors:
                for j in list_neighbors:
                    if i != j and G.has_edge(i, j):
                        links += 1 
            cc_dict[node] = 2 * (links / 2) / (degree * (degree - 1))                
        else:
            cc_dict[node] = 0
            
    return cc_dict

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(clusteringcoefficient(G)==cc_nice(G,n) for n in G)
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
WRITE YOUR CODE HERE

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nicer(G,n) for n in G)
```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    "Geeft aan of een netwerk een volledig netwerk is of niet"
    total_edges = nx.number_of_nodes(G) * (nx.number_of_nodes(G) - 1) / 2
    if nx.number_of_edges(G) == total_edges:
        return True
    else:
        return False
    
is_complete(G)
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```

```python

```