# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


```python
2.1 -
```

```python
2.2 True
```

```python
2.3 False
```

```python
2.4 n-1, het minimum aantal links veranderd als het netwerk niet connected hoeft te zijn, omdat er dan niet tussen elke node een edge hoeft te zijn.
```

```python
2.5 True
```

```python
2.6 True
```

```python
2.7 False
```

```python
# 2.8 Het is een cycle, maar is niet sterk verbonden, wel zwak.

G = nx.DiGraph()

G.add_edges_from([('A', 'B'), ('B', 'C'), ('D', 'B'), ('D', 'E'), ('E', 'F'), ('F', 'A'), ('F', 'B'), ('F', 'C'), ('F', 'D')])

# to convert the graph to undirected and find the connected components:
UG = G.to_undirected()
print(nx.is_strongly_connected(G))
print(nx.is_weakly_connected(G))
print(nx.number_connected_components(UG))
```

```python
# 2.9 Nee

G = nx.Graph()

G.add_edge('A', 'B')
G.add_edge('B', 'C')
G.add_edge('D', 'B')
G.add_edge('D', 'E')
G.add_edge('D', 'F')
G.add_edge('F', 'E')

print(nx.is_tree(G))
```

```python
2.10 4
```

```python
2.11 Het converteren van een zwak verbonden gericht netwerk naar een ongericht netwerk hoeft niet perse te resulteren in een verbonden netwerk, afhankelijk van de oorspronkelijke netwerkstructuur. Als het oorspronkelijke netwerk sterk verbonden is wanneer de richting van randen wordt genegeerd, dan zal het resulterende ongerichte netwerk ook verbonden zijn. Anders zal het ongerichte netwerk niet verbonden zijn.
```

```python
2.12 C
```

```python
# 2.13 C

G = nx.DiGraph()

G.add_edge('A', 'B', weight=2)
G.add_edge('B', 'C', weight=1)
G.add_edge('A', 'C', weight=6)
G.add_edge('C', 'E', weight=3)
G.add_edge('D', 'E', weight=2)
G.add_edge('D', 'F', weight=2)
G.add_edge('B', 'F', weight=11)
G.add_edge('B', 'D', weight=9)

print(G.nodes())
print(G.edges())
```

```python
# 2.14 In-strength is 9, out-strength is 3.

G = nx.DiGraph()

G.add_weighted_edges_from([('A', 'B', 2), ('B', 'C', 1), ('A', 'C', 6),
                           ('C', 'E', 3), ('D', 'E', 2), ('D', 'F', 2),
                           ('B', 'F', 11), ('B', 'D', 9)])

in_strength_D = G.in_degree('D', weight='weight')
print(in_strength_D)

out_strength_C = G.out_degree('C', weight='weight')
print(out_strength_C)
```

```python
2.15 0
```

```python
2.16 B
```

```python
2.17 0.8 + 0.9 + 0.8 + 0.95 = 3.45
```

```python
2.18 3.3725
```

```python
2.19 C
```

```python
2.20 B
```

```python
2.21 Ja
```

```python
2.22 A
```

```python
# 2.23 5

G = nx.DiGraph()

# add edges to the graph
G.add_edges_from([(1,2),(2,5),(5,7),(7,6),(6,3),(3,1),(4,3),(1,4),(4,5),(4,7),(4,6),(2,4)])

# calculate the diameter
diameter = nx.diameter(G)

print(diameter)
```

```python
# 2.24 2

G = nx.Graph()

# add edges to the graph
G.add_edges_from([(1,2),(2,5),(5,7),(7,6),(6,3),(3,1),(4,3),(1,4),(4,5),(4,7),(4,6),(2,4)])

# calculate the diameter
diameter = nx.diameter(G)

print(diameter)
```

```python
2.25 False
```

```python
2.26 C
```

```python
2.27 True
```

```python
2.28 D
```

```python
2.29 -
```

```python
2.30 -
```

```python
2.31 6
```

```python
2.32 1
```

```python
2.33 0
```

```python
2.34 (3 * 1) / 10 = 0.3
```

```python
2.35  
1-2: 1

1-3: 2

1-4: 3

1-5: 4

1-6: 3

1-7: 4

1-8: 6

1-9: 7

1-10: 3

1-11: unreachable

2-3: 1

2-4: 2

2-5: 3

2-6: 2

2-7: 3

2-8: unreachable

2-9: unreachable

2-10: unreachable

2-11: unreachable

3-4: 1

3-5: 2

3-6: 1

3-7: 2

3-8: unreachable

3-9: unreachable

3-10: 1

3-11: unreachable

4-5: 1

4-6: 1

4-7: 1

4-8: unreachable

4-9: unreachable

4-10: 2

4-11: unreachable

5-6: 1

5-7: 2

5-8: unreachable

5-9: unreachable

5-10: 3

5-11: unreachable

6-7: 1

6-8: unreachable

6-9: unreachable

6-10: 2

6-11: unreachable

7-8: unreachable

7-9: unreachable

7-10: 3

7-11: unreachable

8-9: 1

8-10: unreachable

8-11: unreachable

9-10: unreachable

9-11: unreachable

10-11: unreachable
```

```python
2.36 0.278
```

```python
2.37 average is 2.37 
```

```python
2.38 0.348
```

```python
2.39 0.218
```

```python
2.40 A en D
```

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    return max(nx.strongly_connected_components(G), key = len)


def find_OUT(G):
    # WRITE YOUR CODE HERE
    scc = list(find_SCC(G))
    OUT = set()
    for node1 in sorted(G.nodes):
        for node2 in sorted(scc):
            if G.has_edge(node1, node2) == False and G.has_edge(node2, node1) == True and node1 != node2:
                if node1 not in scc or node2 not in scc:
                    OUT.add(node2)
                    # print('added:', [node1, node2])
    
    return OUT

def find_IN(G):
    # WRITE YOUR CODE HERE
    scc = list(find_SCC(G))
    IN = set()
    for node1 in sorted(G.nodes):
        for node2 in sorted(scc):
            if G.has_edge(node1, node2) == True and G.has_edge(node2, node1) == False and node1 != node2:
                if node1 not in scc or node2 not in scc:
                    IN.add(node1)
                    # print('added:', [node1, node2])
    
    return IN

# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
def clusteringcoefficient(G):
    return {node : (len(list(G.neighbors(node))) / (len(G) - 1)) for node in G}

# even er een mooi plotje van maken
pd.Series( clusteringcoefficient(G) ).sort_values().plot(kind='barh',figsize=(8,8))
clusteringcoefficient(G)
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
def cc_nice(G, v):
    cc_dict = dict()
    for node1 in G:
        edges = 0
        for node2 in G:
            if node1 != node2 and G.has_edge(node1, node2):
                edges += 1
        cc_dict[node1] = edges / (len(G) - 1)
    
    return cc_dict

print(cc_nice(G, 0))
# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(clusteringcoefficient(G)==cc_nice(G,n) for n in G)
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
WRITE YOUR CODE HERE

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nicer(G,n) for n in G)
```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    WRITE YOUR CODE HERE
is_complete(G)
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```

```python

```