# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


1) Gedaan

2) True

3) True

4) Een connected netwerk met N nodes kan minimaal N links hebben. Als het netwerk niet connected hoeft te zijn, is het minimum aantal links lager.

5) False

6) True

7) False

8) Er is een cycle tussen f en e in dit netwerk en het netwerk is weakly connected

9) Nee, dit netwerk is geen tree, omdat het nog steeds mogelijk is om bepaalde knooppunten te bereiken nadat bepaalde   verbindingen zijn verwijderd.  Bovendien heeft dit netwerk meer verbindingen dan het minimale aantal vereiste verbindingen voor een boomstructuur (N - 1).

10) De diameter is 5.

11) Nee, het resulterende netwerk hoeft niet verbonden te zijn. Wanneer we een zwak verbonden gericht netwerk converteren naar een ongericht netwerk, negeren we de richting van de verbindingen, maar er kunnen verbindingen zijn die in één richting bestaan maar niet in de andere richting. Als er knooppunten zijn die alleen bereikt kunnen worden door verbindingen in één richting te volgen, dan zullen deze knooppunten niet verbonden zijn met knooppunten die alleen bereikt kunnen worden door verbindingen in de tegenovergestelde richting te volgen, wat resulteert in een niet-verbonden ongericht netwerk.

12) c

13) b

14) De in-sterkte van knooppunt D is 2 en de uit-sterkte van knooppunt C is 1.

15) Er zijn drie knooppunten in de grootste sterk verbonden component.

16) a

17) Pad 1,2,4,6 distance = 2

18) Pad 1,2,4,5,6 distance = 1/3,45

19) C

20) A

21) Een social network zou ongeveer dezelfde diameter hebben als figuur 2.15 maar niet dezelfde clustering coefficient omdat figuur 2.15 nauwelijks driehoek connecties vormt wat wel in een social network zou gebeuren.

22) A

23) 2

24) 2

25) False

26) C

27) True

28) D

29) William Timothy Gowers > Peter Clive Sarnak = 4, Martin Hairer > Bruce C. Baguley = 4

30) > Adam Sandler  = 4, Li Xian = 3

31) wordt niet opgeslagen, dus kan je ook niet zien. Kwam wel ver bij de tweede poging.

32) 1


33) 0

34) 3 (edges tussen neighbors) / 15 (mogelijke edges tussen neighbors) = 0.2

35) the shortest path lengths in figure 2.4 differ from 1, 2, 3 and 4

36) the clustering coefficient of these nodes in figure 2.4 differ from 0, 0.167, 0.25, 0.33 to 1

37) the shortest path lengths in figure 2.12 differ from 1, 2, 3 and 4, the average shortest path length of the network is (a=2.2857 + f=2.2857 + g=1.5714 + c=1.5714 + b=1.4286 + h=2.2857 + e=1.8571 + d=2.7143) / 8 = 2

38) a = 1, b = 0.33, c = 0.66, e = 0.33, f = 1, g = 0. 33, average = (1 + 0.33 + 0.66 + 0.33 + 1 + 0.33) = 0.61

39) laten we zeggen dat ik 600 volgers heb, waarvan er 300 onderling elkaar volgen. Dan is het mogelijke aantal links 600 x (600-1)/2 = 179,700. Dan is C = 2 x 300/179,700 = 0.00105

40) a, d

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    return max(nx.strongly_connected_components(G), key=len)

def find_OUT(G):
    scc = find_SCC(G)
    return {node for node in G if node not in scc for scc_node in scc
            if node in G.successors(scc_node) and node not in G.predecessors(scc_node)}

def find_IN(G):
    scc = find_SCC(G)
    return {node for node in G if node not in scc for scc_node in scc
            if node in G.predecessors(scc_node) and node not in G.successors(scc_node)}

# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
import math

def clusteringcoefficient(G):
    clusco = {}

    for node in G.nodes():
        neighbors = list(G.neighbors(node))
        neighbor_count = len(neighbors)
        if neighbor_count <= 1:
            clusco[node] = 0.0
        else:
            triangles = 0
            for x in range(neighbor_count):
                for y in range(x+1, neighbor_count):
                    if G.has_edge(neighbors[x], neighbors[y]):
                        triangles += 1
            clusco[node] = 2 * triangles / (neighbor_count * (neighbor_count - 1))
    return clusco
        
    
# even er een mooi plotje van maken
pd.Series( clusteringcoefficient(G) ).sort_values().plot(kind='barh',figsize=(8,8))

clusteringcoefficient(G)
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
def cc_nice(G, v):
    neighbors = list(G.neighbors(v))
    neighbor_count = len(neighbors)
    if neighbor_count <= 1:
        return 0.0
    else:
        triangles = 0
        for x in range(neighbor_count):
            for y in range(x+1, neighbor_count):
                if G.has_edge(neighbors[x], neighbors[y]):
                    triangles += 1
        return 2 * triangles / (neighbor_count * (neighbor_count - 1))


# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nice(G,n) for n in G)
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
# Om de CC van v te achterhalen, moet de density van de buren van v worden berekend

def cc_nicer(G, v):
    
    #zet buren van v in een lijst
    neighbors_of_v = list(G.neighbors(v)) 

    # creeer subgraph met buren van v
    neighbors_subgraph = G.subgraph(neighbors_of_v)
    
    # bereken density van subgraph
    clustering_coefficient = nx.density(neighbors_of_v)

    return clustering_coefficient

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nicer(G,n) for n in G)

```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    if (G.number_of_nodes()*(G.number_of_nodes()-1))/2 == G.number_of_edges():
        return True
    else:
        return False

is_complete(G)
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```