# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 3
* Alle vragen zijn één punt waard

```python
%matplotlib inline
import networkx as nx
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
```

```python

```

# Deel 1: First Course Chap 2

* Maak vragen 1-40

### Ook met programmeren

* Check je antwoorden op vragen 8-11 door het gerichte netwerk te programmeren, en alle vragen ook te programmeren in Networkx.
* Zet de netwerken in Fig 2-13 en 2.14 in Networkx, probeer ze ook zo mooi uit te printen (zie assignment week 2), en check je antwoorden op de vragen erover.
* Als je het leuk vindt kan je dit ook met de andere voorbeelden doen. Je weet dan wel zeker dat je het goed hebt.
    * Op het tentamen mag je ook programmeren.....


```python
# GEMAAKT VOOR BOEKOPDRACHTEN
print('')
print('Vraag 2.8 — gewogen en ongericht')

D = nx.DiGraph()

D.add_nodes_from(['A', 'B', 'C', 'D', 'E', 'F'])

D.add_weighted_edges_from([('A', 'B', 1), ('B', 'C', 2), ('D', 'B', 1), ('D', 'E', 1), ('E', 'F', 1), ('F', 'A', 2), ('F', 'B', 1), ('F', 'C', 3), ('F', 'D', 1), ('F', 'E', 1)])

pos = nx.spring_layout(D)  # Geeft de posities van de knopen
nx.draw_networkx_nodes(D, pos)  # Tekent de knopen
nx.draw_networkx_edges(D, pos, width=1)  # Tekent de verbindingen
nx.draw_networkx_labels(D, pos)  # Tekent de labels

# Voegt verbindingslabel toe aan de hand van de gewichten
labels = nx.get_edge_attributes(D, 'weight')
nx.draw_networkx_edge_labels(D, pos, edge_labels=labels)

plt.axis('off')
plt.show()

# Is het netwerk sterk verbonden?
print(nx.is_strongly_connected(D))

###############################################################
print('')
print('Vraag 2.9 en 2.10 — niet-gewogen en ongericht')

G = nx.Graph()

G.add_nodes_from(['A', 'B', 'C', 'D', 'E', 'F'])

G.add_edges_from([('A', 'B'), ('B', 'C'), ('D', 'B'), ('D', 'E'), ('E', 'F'), ('F', 'A'), ('F', 'B'), ('F', 'C'), ('F', 'D'), ('F', 'E')])

nx.draw(G, with_labels=True)

plt.axis('off')
plt.show()

# Is het netwerk sterk verbonden?
print(nx.is_tree(G))
print(nx.diameter(G))
print(nx.is_connected(G))


###############################################################
print('')
print('Vraag 2.13, 2.14 en 2.15')

D = nx.DiGraph()

D.add_nodes_from(['A', 'B', 'C', 'D', 'E', 'F'])

D.add_weighted_edges_from([('B', 'F', 11), ('B', 'D', 9), ('D', 'F', 2), ('E', 'D', 2), ('C', 'E', 3), ('A', 'C', 6), ('B', 'C', 1), ('A', 'B', 2)])

pos = nx.spring_layout(D)  # Geeft de posities van de knopen
nx.draw_networkx_nodes(D, pos)  # Tekent de knopen
nx.draw_networkx_edges(D, pos, width=1)  # Tekent de verbindingen
nx.draw_networkx_labels(D, pos)  # Tekent de labels

# Voegt verbindingslabel toe aan de hand van de gewichten
labels = nx.get_edge_attributes(D, 'weight')
nx.draw_networkx_edge_labels(D, pos, edge_labels=labels)

plt.axis('off')
plt.show()

# Vindt het sterkst verbonden component, en geeft zijn knopenaantal
sterkst_verbonden_componenten = max(nx.strongly_connected_components(D), key = len) # F
aantal_knopen_sterkste = len(sterkst_verbonden_componenten)


# Is het netwerk sterk verbonden?
print("Dit netwerk is sterk verbonden:", nx.is_strongly_connected(D))
print("Dit netwerk is zwak verbonden:", nx.is_weakly_connected(D))

print("Knoop D's inkracht is:", D.in_degree('D', weight='weight'))
print("Knoop C's uitkracht is:", D.out_degree('C', weight='weight'))
print("Aantal knopen in sterkst verbonden component:", aantal_knopen_sterkste)

###############################################################
print('')
print('Vraag 2.16, 2.17 en 2.18')

D = nx.DiGraph()

D.add_nodes_from(['1', '2', '3', '4', '5', '6'])

D.add_weighted_edges_from([('1', '2', 0.8), ('2', '4', 0.9), ('4', '6', 0.3), ('4', '5', 0.8), ('3', '5', 0.5), ('1', '3', 0.6), ('5', '6', 0.95)])

pos = nx.spring_layout(D)  # Geeft de posities van de knopen
nx.draw_networkx_nodes(D, pos)  # Tekent de knopen
nx.draw_networkx_edges(D, pos, width=1)  # Tekent de verbindingen
nx.draw_networkx_labels(D, pos)  # Tekent de labels

# Voegt verbindingslabel toe aan de hand van de gewichten
labels = nx.get_edge_attributes(D, 'weight')
nx.draw_networkx_edge_labels(D, pos, edge_labels=labels)

plt.axis('off')
plt.show()

# Is het netwerk sterk verbonden?
print("Dit netwerk is sterk verbonden:", nx.is_strongly_connected(D))
print("Dit netwerk is zwak verbonden:", nx.is_weakly_connected(D))
print("De korste route van 1 naar 6:", nx.shortest_path(D, '1', '6'))
print("De korste gewogen route van 1 naar 6:", nx.shortest_path(D, '1', '6', weight='weight'))


```

```python

###############################################################
print('')
print('Vraag 2.22 en 2.23')

D = nx.DiGraph()

D.add_nodes_from(['1', '2', '3', '4', '5', '6', '7'])

D.add_edges_from([('1', '2'), ('1', '4'), ('2', '4'), ('2', '5'), ('3', '1'), ('4', '3'), ('4', '6'), ('4', '7'), ('4', '5'), ('5', '7'), ('6', '3'), ('7', '6')])

# Tekent het netwerk
nx.draw(D, with_labels=True)

# Is het netwerk sterk verbonden of verbonden, etc.?
print("Dit netwerk is sterk verbonden:", nx.is_strongly_connected(D))
print("Dit netwerk is zwak verbonden:", nx.is_weakly_connected(D))
print("De diameter van dit netwerk is:", nx.diameter(D)) # 5

###############################################################
print('')
print('Vraag 2.24')

G = nx.Graph()

G.add_nodes_from(['1', '2', '3', '4', '5', '6', '7'])

G.add_edges_from([('1', '2'), ('1', '4'), ('2', '4'), ('2', '5'), ('3', '1'), ('4', '3'), ('4', '6'), ('4', '7'), ('4', '5'), ('5', '7'), ('6', '3'), ('7', '6')])

# Tekent het netwerk
nx.draw(G, with_labels=True)

print("De diameter van dit netwerk is:", nx.diameter(G)) # 2

###############################################################
print('')
print('Vraag 2.35 en 2.36')

G = nx.Graph()

G.add_nodes_from(['1', '2', '3', '4', '5', '6', '7', '8', '9', '10', '11'])

G.add_edges_from([('1', '2'), ('2', '4'), ('2', '3'), ('4', '3'), ('3', '6'), ('3', '5'), ('6', '8'), ('6', '7'), ('10', '11')])

# tekent het netwerk
nx.draw(G, with_labels=True)

# berekent de kortste padlengte voor elk paar knooppunten in het gigantcomponent
kortste_padlengten = dict(nx.shortest_path_length(G))
print(kortste_padlengten)

#print('')
print('De korste padlengten tussen elk knopenpaar:', dict(nx.shortest_path_length(G)))
        
# print de clustervormingscoëfficiënt voor iedere knoop
print("De clustervormingscoëfficiënten:", nx.clustering(G))

```

```python
###############################################################
print('')
print('Vraag 2.37 en 2.38')

G = nx.Graph()

G.add_nodes_from(['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h'])

G.add_edges_from([('a', 'g'), ('a', 'f'), ('b', 'c'), ('b', 'e'), ('b', 'g'), ('b', 'h'), ('c', 'e'), ('c', 'g'), ('d', 'e'), ('f', 'g')])

# tekent het netwerk
nx.draw(G, with_labels=True)

# berekent de kortste padlengte voor elk paar knooppunten in het gigantcomponent
print('')
print('De korste padlengten tussen elk knopenpaar:', dict(nx.shortest_path_length(G)))

# Berekent de gemiddelde kortstepadlengte van het netwerk
print('')
print("De gemiddelde kortstepadlengte van het netwerk is:", nx.average_shortest_path_length(G)) # 2 

# print de clustervormingscoëfficiënt voor iedere knoop
print('')
print("De clustervormingscoëfficiënten van alle knopen:", nx.clustering(G))

# print de clustervormingscoëfficiënt voor het netwerk
print('')
print("De clustervormingscoëfficiënt van het netwerk is:", nx.average_clustering(G))
```

# Deel 2: Easley & Kleinberg Chapter 13

De volgende vragen hebben betrekking op [hoofdstuk 13 van Easley and Kleinberg](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch13.pdf). Hieronder staat de networkx graaf gebaseerd op figuur 13.8, pas deze vooral aan om te zorgen dat je functies altijd werken.

```python
EK = nx.DiGraph()
EK.add_edges_from([(6, 2), (6, 7), (7, 8), (11, 7), (1, 8), (12, 13), (8, 13),
                   (18, 13), (15, 18), (15, 14), (13, 14), (14, 9), (9, 3),
                   (3, 8), (4, 1), (4, 5), (5, 10), (16, 10), (17, 16), (15,
                                                                         16),
                   (9, 4), (9, 15)])

nx.draw_circular(EK, with_labels=True)
```

## Vraag 5

* Maak een functie `find_SCC` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de grootste SCC (strongly connected component).

* Maak een functie `find_IN` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de IN (dat zijn alle knopen met een pad naar de SCC, maar vanuit de SCC kan je niet naar die knopen komen).

* Maak een functie `find_OUT` die gegeven een gerichte graaf een set teruggeeft met daarin alle nodes die behoren tot de OUT (ja, dit snap je nu zelf wel hoop ik).

<!-- grade_id: Vraag5_antwoord, points: 0, grade: false, solution: true -->
```python
def find_SCC(G):
    # bepaalt alle sterk verbonden componenten (scc of svc)
    scc_list = [scc for scc in nx.strongly_connected_components(G)]
    # bepaal de grootte van de grootste
    max_scc_size = max(len(scc) for scc in scc_list)
    # bepaalt de knopen van de svc
    return [scc for scc in scc_list if len(scc) == max_scc_size][0]

def find_OUT(G):
    scc = find_SCC(G)
    # voegt alle inkomende buren van knopen in de svc toe aan de set inknopen, 
    # en verwijdert de knopen die er al in zitten
    return {node for n in scc for node in G.successors(n)} - scc

def find_IN(G):
    scc = find_SCC(G)
    # doet het omgekeerde van find_OUT(G)
    return {node for n in scc for node in G.predecessors(n)} - scc


# even testen op figuur 13.8 uit het boek
find_SCC(EK), find_IN(EK), find_OUT(EK)
```

<!-- grade_id: cell-dfc21dad8b2e3906, points: 1, grade: true, solution: false -->
```python
assert_equal(type(find_SCC(EK)), set)
assert_equal(type(find_OUT(EK)), set)
assert_equal(type(find_IN(EK)), set)
assert find_SCC(EK).issubset(EK.nodes())
assert find_IN(EK).issubset(EK.nodes())
assert find_OUT(EK).issubset(EK.nodes())
```

```python
# resultaat graaf tekenen
def draw_web(G):
    pos = nx.spring_layout(G)
    nx.draw(G, pos, with_labels=True)
    color = sns.color_palette("hls", 7)
    #tubes = find_tubes(G)
    legenda = [
        'SCC', 'IN', 'OUT', ] #'Tendrils (IN)', 'Tendrils (OUT)', 'Tubes','Disconnected' ]
    listofnodes = [
        find_SCC(G),
        find_IN(G),
        find_OUT(G),
        #find_tendrils_IN(G) - tubes,
        #find_tendrils_OUT(G) - tubes, tubes,
        #find_disconnected(G)
    ]
    for n, (label, nodes, c) in enumerate(zip(legenda, listofnodes, color)):
        nx.draw_networkx_nodes(
            G, pos, nodelist=nodes, node_color=[c], label=label)
    nx.draw_networkx_edges(G, pos)
    plt.legend()
    plt.axis('off')
    plt.show()


draw_web(EK)
```

# Deel 3: Clustering coëfficiënt

```python
G = nx.karate_club_graph()

nx.draw(G, with_labels=True)
```

## Vraag 3.1
Maak een functie `clusteringcoefficient(G)` die een graaf accepteert als input, en voor alle knopen de clusteringcoëfficiënt ([Sectie 3.1 E&amp;K](http://www.cs.cornell.edu/home/kleinber/networks-book/networks-book-ch03.pdf)) berekent.
                
<strong>Bereken dit zelf, gebruik geen NetworkX functie die dit al doet zoals `clustering` en `average_clustering`. </strong> Je mag natuurlijk wel NetworkX handigheidjes als `G.neighbors(v)` gebruiken.

Als output krijg je een dict  met de knopen van `G` als sleutels en de clustering coeffiecient van die knoop als bijbehorende waarde.   Geef waarde 0 als een knoop degree 1 heeft.


De verwachte output voor deze graaf is 1 voor knoop 7 en 2/3 voor knoop 3 en 0 voor knoop 11. etc. 
  

<!-- grade_id: Vraag7, points: 0, grade: false, solution: true -->
```python
from itertools import combinations

def clusteringcoefficient(G):
    return {node: 0 if G.degree(node) == 1 else sum(1 for u, v in combinations(G.neighbors(node), 2) \
           if G.has_edge(u, v)) / (G.degree(node) * (G.degree(node) - 1) / 2) for node in G.nodes()}


# even er een mooi plotje van maken
pd.Series( clusteringcoefficient(G) ).sort_values().plot(kind='barh',figsize=(8,8))

clusteringcoefficient(G)
```

<!-- grade_id: Vraag7_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(clusteringcoefficient(G), dict)
for n in clusteringcoefficient(G):
    assert n in G
    assert isinstance(clusteringcoefficient(G)[n],float)
```

## Vraag 3.1.1

Maak de functie `cc_nice(G, v)` nu nog een keer maar gebruik nu alleen maar `G.has_edge()` van networkx.
Chack dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 311a, points: 1, grade: true, solution: true -->
```python
def cc_nice(G, v=None):
    cc_dict = {}
    for node in G.nodes():
        neighbors = set(G.neighbors(node))
        k = len(neighbors)
        if k < 2:
            cc_dict[node] = 0
            continue
        total = 0
        for neighbor in neighbors:
            for neighbor2 in neighbors:
                if neighbor == neighbor2:
                    continue
                if G.has_edge(neighbor, neighbor2):
                    total += 1
        cc_dict[node] = total / (k * (k - 1))
    if v is not None:
        return cc_dict[v]
    else:
        return cc_dict
    

def cc(G,n): # networkx's functie            
    return nx.clustering(G)

# de eerste functie
print('')
print("De functie van vraag 3.1:", clusteringcoefficient(G))

# networkx-functie
print('')
print("networkx-functie:", cc(G, n))

# de nieuwste functie
print('')
print("Nieuwe functie:", cc_nice(G))


# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nice(G) for n in G)
```

## Vraag 3.1.2

Maak de functie `cc_nicer(G, v)` nu nog een keer maar gebruik nu alleen maar  `G.subgraph` en `nx.density` van networkx.
Check dat hij echt hetzelfde doet als je eerdere functie en als networkx clustering.

<!-- grade_id: 312a, points: 1, grade: true, solution: true -->
```python
def cc_nicer(G):
    return {n: nx.density(G.subgraph(set(G.neighbors(n)))) if G.degree(n) > 1 else 0 for n in G.nodes()}


# de nieuwste functie
print('')
print("Nieuwe functie:", cc_nicer(G))

# check dat de twee functies hetzelfde resultaat geven voor elke knoop op de karateclub
all(cc(G,n)==cc_nicer(G) for n in G)
```

## Vraag 3.2

2. Maak de functie  `is_complete(G)`  die `True` teruggeeft precies dan als het ongerichte netwerk  `G` een volledig netwerk is en `False` als dat niet het geval is.  Gebruik **alleen** de networkx functies `nx.number_of_nodes(G)` en `G.number_of_edges()`.

<!-- grade_id: Vraag8, points: 0, grade: false, solution: true -->
```python
def is_complete(G):
    # Lmax = N(N - 1) / 2
    n = nx.number_of_nodes(G)
    l_max = nx.number_of_edges(G)
    return l_max == ((n * (n - 1)) / 2) # bool() is niet nodig
    
is_complete(G)
```

<!-- grade_id: Vraag8_antwoord, points: 1, grade: true, solution: false -->
```python
assert isinstance(is_complete(G), bool)
```

```python

```

```python

```